Testosterone-Induced Metabolic Changes in Seminal Vesicle Epithelium Modify Seminal Plasma Components with Potential to Improve Sperm Motility

Author response:

The following is the authors’ response to the previous reviews

Editor's note:

Thank you for taking time and efforts to improve this study. After re-review, two reviewers have a consensus that the connections the fatty acids and sperm motility is still ambiguous. Thus, I recommend to further tone down this conclusion consistently in the title and the text pointed out by reviewers before making a final version of record.

We sincerely appreciate the considerable time and effort you and the reviewers devoted to evaluating our manuscript. We have revised the title and text to express the relationship between fatty acids and sperm motility more consistently and toned down. With these revisions, we would like to proceed with publishing the manuscript as the Version of Record (VoR). Thank you very much for your guidance in improving our study.

Public Reviews:

Reviewer #1 (Public review):

Summary:

In this revised report, Yamanaka and colleagues investigate a proposed mechanism by which testosterone modulates seminal plasma metabolites in mice. Based on limited evidence in previous versions of the report, the authors softened the claim that oleic acid derived from seminal vesicle epithelium strongly affects linear progressive motility in isolated cauda epididymal sperm in vitro. Though the report still contains somewhat ambiguous references to the strength of the relationship between fatty acids and sperm motility.

Strengths:

Often, reported epidydimal sperm from mice have lower percent progressive motility compared with sperm retrieved from the uterus or by comparison with human ejaculated sperm. The findings in this report may improve in vitro conditions to overcome this problem, as well as add important physiological context to the role of reproductive tract glandular secretions in modulating sperm behaviors. The strongest observations are related to the sensitivity of seminal vesicle epithelial cells to testosterone. The revisions include the addition of methodological detail, modified language to reflect the nuance of some of the measurements, as well as re-performed experiments with more appropriate control groups. The findings are likely to be of general interest to the field by providing context for follow-on studies regarding the relationship between fatty acid beta oxidation and sperm motility pattern.

Weaknesses:

The connection between media fatty acids and sperm motility pattern remains inconclusive.

We would like to express our sincere gratitude to the judges for their cooperation in reviewing the manuscript and for your helpful comments, which were instrumental in improving manuscript.

Reviewer #2 (Public review):

Using a combination of in vivo studies with testosterone-inhibited and aged mice with lower testosterone levels as well as isolated mouse and human seminal vesicle epithelial cells the authors show that testosterone induces an increase in glucose uptake. They find that testosterone induces a difference in gene expression with a focus on metabolic enzymes. Specifically, they identify increased expression of enzymes regulating cholesterol and fatty acid synthesis, leading to increased production of 18:1 oleic acid. The revised version strengthens the role of ACLY as the main regulator of seminal vesicle epithelial cell metabolic programming. The authors propose that fatty acids are secreted by seminal vesicle epithelial cells and are taken up by sperm, positively affecting sperm function. A lipid mixture mimicking the lipids secreted by seminal vesicle epithelial cells, however, only has a small and mostly non-significant effect on sperm motility, suggesting the authors were not apply to pinpoint the seminal vesicle fluid component that positively affects sperm function.

We greatly appreciate the reviewer’s thoughtful comments and time spent reviewing this manuscript. The relationship between lipids such as fatty acids and sperm motility remains unclear in the current dataset. Therefore, before finalizing the manuscript, we revised the title and text, as suggested by the reviewers, to express this conclusion more cautiously and consistently.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Some additional comments are provided below to aid the authors in improving the quality of the work:

Major Comments:

(1) In the newly added supplemental figure 5, the authors note that the percentage data were arcisine transformed prior to statistical analysis without providing any other justification. This seems strange, especially for such a small sample size. It seems more appropriate for the authors to use a nonparametric test. Forcing symmetry without knowing what the shape of the true distribution is makes the ANOVA hard to interpret. Additionally, why use pairwise comparisons rather than comparing each group to the control (LM 0%). Also, note that the graphs are not individually labeled to distinguish them in the legend (A, B, C, etc.). Ultimately, the treatment differences don't seem that meaningful, even if the authors were able to observe statistical significance with the somewhat over-manipulated method of analysis.

Ultimately, the conclusion of this experiment (Supplemental figure 5) remains unchanged, but we agree that the relationship between fatty acids and sperm motility remains unclear. Therefore, before finalizing the manuscript, we revised the title and text as pointed out by the reviewers to express this conclusion more cautiously and consistently throughout the manuscript.

Arcsin transform is commonly used for percentage data [Zar, J.H. 2010. Biostatistical analysis., McDonald, J.H. 2014. Handbook of biological statistics.]. If the values are low or high, such as 0 to 30% or 70 to 100%, without arcsine transformation will result in a large deviation from the normality of the data. However, even if such a conversion is performed, it does not necessarily mean that the assumptions of normality and homogeneity of variance, which are prerequisites for parametric statistical analysis methods, are satisfied.

Given the small sample size and the possibility of non-normal data, we performed Shapiro–Wilk tests for each group (n = 6) and found no departure from normality (all p > 0.1). Q–Q plots and Levene’s test (p > 0.1) likewise supported the assumptions of ANOVA. Following the reviewer’s recommendation, we repeated the analysis with a Kruskal–Wallis test followed by Dunn’s post-hoc comparisons (Bonferroni corrected). Both approaches led to the same conclusions, with non-parametric p-values equal to or smaller than the parametric ones. In the revised manuscript we now report ANOVA as the primary analysis. The author response image includes effect sizes with 95 % confidence intervals, and provide the non-parametric results for transparency.

Results of reanalysis of supplementary Figure 5 using nonparametric tests and effect sizes with 95% confidence intervals.

Upper part; Differences between groups were assessed by Kruskal–Wallis test, differences among values were analyzed by Dunn’s post-hoc comparisons (Bonferroni corrected) for multiple comparisons. Different letters represent significantly different groups. Lower part; The effect sizes with 95 % confidence intervals. For example, Cliff's Δ = -1 (95% CI ~ -0.6) in VSL's “LM 0 vs LM1” means that LM 1% values exceed LM 0 %values in all pairs.

(2) I appreciate that the authors toned down the interpretation of the effects of seminal plasma metabolites on sperm motility with a cautionary statement on Lines 397-405 and Line 259. However, they send mixed signals with the title of the report: "Testosterone-Induced Metabolic Changes in Seminal Vesicle Epithelial cells Alter Plasma Components to Enhance Sperm Motility", and on line 265 when the say "ACLY expression is upregulated by testosterone and is essential for the metabolic shift of seminal vesicle epithelial cells that mediates sperm linear motility".

The wording has been softened overall. The title has been changed to “Testosterone-Induced Metabolic Changes in Seminal Vesicle Epithelium Modify Seminal Plasma Components with Potential to Improve Sperm Motility” In the results (lines 265-266), we have stated that “ACLY expression is upregulated by testosterone and is essential for the metabolic shift that is associated with increased linear motility” without implying a causal relationship.

Minor Comments:

(1) Typo on line 31: "understanding the male fertility mechanisms and may perspective for the development of potential biomarkers of male fertility and advance in the treatment of male infertility."

We have made the following corrections. “These findings suggest that testosterone-dependent lipid remodeling may contribute to sperm straight-line motility, and further functional verification is required.”

(2) Line 193: the statement is confusing "Therefore, we analyzed mitochondrial metabolism using a flux analyzer, predicting that more glucose is metabolized, pyruvate is metabolized from phosphoenolpyruvic acid through glycolysis in response to testosterone, and is further metabolized in the mitochondria." For example, 'Metabolized through glycolysis' is an ambiguous way to describe the pyruvate kinase reaction. Additionally, phosphoenolpyruvate has three acid ionizable groups, two of which have pKa's well below physiological pH, so phosphoenolpyruvate is the correct intermediate rather than phosphoenolpyruvic acid. The authors make similar mistakes with other organic acids such as citric acid.

Rewritten as “We therefore examined cellular energy metabolism with a flux analyzer, anticipating that testosterone would elevate glycolytic flux, thereby producing more pyruvate from phosphoenolpyruvate. Because extracellular pyruvate levels simultaneously declined, we inferred that the cells had an increased pyruvate demand and, at that time, hypothesized that the excess pyruvate would enter the mitochondria to support enhanced oxidative metabolism.” (lines 193-198)

The organic acids are now referenced in their appropriate forms (e.g., citrate, phosphoenolpyruvate).

(3) Line: 271: "Acly" should be all capitalized to "ACLY". The report mixes capitalizing through out and could be more consistent.

We appreciate the reviewer’s attention to nomenclature and have standardized the manuscript accordingly. Proteins are written in Roman letters, all in capital letters. Mouse gene symbols: italics, first letter capitalize.

Reviewer #2 (Recommendations for the authors):

Major comments:

(1) 'Once capacitation is complete, sperm cannot maintain that state for a long time'. The publications cited by the author do not support that statement and this reviewer also does not agree. Lower fertilization efficiency from in vitro capacitated epidydimal sperm does not have to mean capacitation is reversed, it can simply mean in vitro capacitation conditions not accurately mimic capacitation in vivo.

We thank the reviewer for pointing this out and would like to clarify our position. Our statement does not suggest a "reversal" of active capacitation. Rather, it reflects the well-documented fact that capacitation is a transient process. Sperm that undergo capacitation too early cannot maintain that state for long enough to retain their ability to fertilize at the moment and location of fertilization in vivo.

(2) How do the authors explain the discrepancy between the results shown in Fig. S1E, the increase in sperm motility upon mixing of sperm with SVF and the results reported in Li et al 2025. Mentioning decapacitating factors without further explanation is insufficient.

We appreciate the reviewer's feedback pointing out the need for a clearer explanation.

Seminal plasma is inherently binary, containing both decapacitation factors that delay or inhibit capacitation and nutrient substrates that promote sperm motility.

In vivo, it is believed that the coating of sperm by decapacitation factors is removed by uterine fluid and albumin as it passes through the female reproductive tract [PMID: 22827391, PMID: 24274412]. In contrast, standard fertilization culture media lack a clearance pathway, so decapacitating factors are retained throughout the culture period. As a result, the cleavage rate after in vitro fertilization using sperm exposed to seminal vesicle fluid decreased dramatically.

Lipids, such as fatty acids, increased sperm motility without directly inducing markers of fertilization. These results suggest that the enhancement of motility by lipids is functionally distinct from the capacitation-inhibiting function of seminal plasma proteins. The data from this study are consistent with the biphasic model. Specifically, decapacitation factors temporarily stabilize the sperm membrane, preventing early capacitation. Meanwhile, lipids enhance sperm motility, enabling them to rapidly pass through the hostile uterine environment.

(3) This reviewer does not see the merit in including a lipid mixture motility experiment compared to using OA alone. The increase in motility is still small and far from comparable to the motility increase with seminal vesicle fluid. In this reviewer's opinion the experiment is still inconclusive and should not be highlighted in the manuscript title.

The wording has been softened overall. The title has been changed to “Testosterone-Induced Metabolic Changes in Seminal Vesicle Epithelium Modify Seminal Plasma Components with Potential to Improve Sperm Motility”. (Please see also Reviewer 1's main comment 1)

Minor comments:

(1) 'This change includes a large amplitude of flagella' does not make sense. Please correct.

The following corrections have been made. “This change is characterized by large-amplitude flagellar beating.” (lines 44-45)

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Male factors account for almost half of the causes of infertility. In rodents and humans, most of the components of semen are supplied by the seminal vesicles, and they support male reproductive ability, but there are many unknown details. In this study, the metabolic changes in the seminal vesicle epithelial cells were focused on, and the mechanisms by which testosterone affects the seminal plasma composition were investigated. A factor that improves the linear motility of sperm was secreted from the seminal vesicles, which were synthesized in an androgen-dependent manner. Bioassays, gene expression, and flux analysis studies demonstrated that testosterone promotes glucose uptake in seminal vesicle epithelial cells via GLUT4, resulting in oleic acid synthesis. Oleic acid was shown to be taken up by sperm and to promote linear motility, thereby improving fertilization rates both in vitro and in vivo. ACLY was a critical factor in this metabolic change, which produces oleic acid and enhances their fertilization ability in vivo. In conclusion, the critical role of testosterone-induced metabolic changes in the seminal vesicles is to ensure the synthesis of oleic acid, which is important for sperm fertilization in vivo. These findings provide new perspectives for the development of potential biomarkers of male fertility and advances in the treatment of male infertility.

Revised version: Male factors account for almost half of the causes of infertility. In rodents and humans, most of the components of semen are supplied by the seminal vesicles, and they support male reproductive ability, but there are many unknown details. In this study, the metabolic changes in the seminal vesicle epithelial cells were focused on, and the mechanisms by which testosterone affected the seminal plasma composition were investigated. A factor improving the linear motility of sperm was secreted from the seminal vesicles, and it was produced in an androgen-dependent manner. Bioassays, gene expression, and flux analysis studies demonstrated that testosterone promotes glucose uptake in seminal vesicle epithelial cells via GLUT4, resulting in fatty acid synthesis. ACLY was a critical factor in this metabolic change, which produced fatty acid, especially oleic acid. In conclusion, the critical role of testosterone-induced metabolic changes in the seminal vesicles is to ensure the synthesis of fatty acid. These findings suggest that testosterone-dependent lipid remodeling may contribute to sperm straight-line motility, and further functional verification is required.

Introduction

Earlier version: Infertility is defined by the World Health Organization (WHO) as the failure to conceive within 12 months despite contraceptive-free sexual intercourse. Its prevalence has increased over the decades and accounts for 17.5% of couples worldwide (1). The male partner is thought to be solely or in combination responsible for 50% of infertility cases, and the male factor in the etiology of infertility is highly recognized. Male fertility depends largely on the quality of semen composed of sperm and seminal plasma and decreases with aging(2–4).

Revised version: Infertility is defined by the World Health Organization (WHO) as the failure to conceive within 12 months despite contraceptive-free sexual intercourse. Its prevalence has increased over the decades and is observed in 17.5% of couples worldwide (1). The male partner is thought to be solely or in combination responsible for 50% of infertility cases, and the male factor in the etiology of infertility is highly recognized. Male fertility depends largely on the quality of semen composed of sperm and seminal plasma and decreases with aging(2–4).

Added in the revised version: Sperm that have completed spermatogenesis in the testis acquire the potential to fertilize during maturation in the epididymis (5–7). The physiological change of sperm during fertilization process are collectively referred to as “capacitation”. This change is characterized by large-amplitude flagellar beating (called hyperactivation) and developing the capacity to undergo the acrosome reaction, and can be induced by culturing sperm collected from the epididymis in a medium containing bicarbonate and cholesterol acceptors (8, 9). However, once capacitation is complete, sperm cannot maintain this state for a long time. Therefore, even if epididymal sperm that have not been exposed to seminal plasma are artificially inseminated into the cervix or uterus, the fertilization rate remains low (10–12). This is because, in vivo, during ejaculation, exposure of epididymal sperm to seminal plasma masks the unintended capacitation as they pass through the female reproductive tract and ensures fertilization of sperm that reaches the oviduct (13). In other words, seminal plasma plays an important role in fine-tuning the timing of sperm capacitation and in maintaining the sustained sperm motility required to reach the oviduct.

Removed in the revised version: Spermatogenesis occurs in the testes. While passing through the epididymis, sperm undergo a maturation process, but sperm do not fully acquire their fertilizing ability in the epididymis. (5–7). Therefore, sperm collected from the epididymis cannot fertilize oocytes unless they acquire the ability to fertilize in in vitro and cannot maintain sustained liner motility required for migration from the uterus to the oviduct (8, 9). Thus, when epididymal sperm are artificially inseminated to the cervix or uterus, fertilization rates are low (10–12). However, when epididymal sperm are mixed with seminal plasma at the time of ejaculation or when sperm are exposed to seminal plasma in vitro, even for only a brief period, they acquire in vivo fertilization potential and are capable of passing through the female reproductive tract and fertilizing ovulated oocyte in vivo (13). Thus, sperm formation occurs in the testis; they mature in the epididymis but only acquire the in vivo fertilization ability after exposure to seminal plasma.

Earlier version: Seminal plasma is mainly produced in the seminal vesicles (approximately 65-75%), with approximately 20-30% produced by the prostate, and small amounts secreted by the bulbourethral and urethral glands (14). The mixing of seminal fluid with accessory reproductive gland fluid during ejaculation causes semen to clot and liquefy in humans and causes vaginal plug formation in rodents. Vaginal plugs not only inhibit continuous mating but also play a role in maintaining sperm in the uterus and ensuring male fertility (13). TGFβ, a component of seminal plasma, increases antigen-specific Treg cells in the uterus of mice and humans, which induces immune tolerance, resulting in pregnancy (15, 16). Furthermore, SVS2, secreted from the seminal vesicle, is also known to be required for sperm to escape attack by the female immune system in the uterus (17). In addition, seminal plasma contains various biochemical components that support sperm metabolism, including lipids and organic acids. Differences in the relative size of the accessory reproductive glands reflect differences in the seminal plasma composition. Accordingly, there are considerable differences in the properties of seminal plasma between animal species. Specifically,

Revised version: Seminal plasma is mainly produced in the seminal vesicles (approximately 65-75%), with approximately 20-30% produced by the prostate, and small amounts secreted by the bulbourethral and urethral glands (14). The mixing of seminal fluid with accessory reproductive gland fluid during ejaculation causes semen to clot and liquefy in humans and causes vaginal plug formation in rodents. Vaginal plugs not only inhibit continuous mating but also play a role in maintaining sperm in the uterus and ensuring male fertility (13). TGFβ, a component of seminal plasma, increases antigen-specific Treg cells in the uterus of mice and humans, which induces immune tolerance, resulting in pregnancy (15, 16). Furthermore, SVS2, secreted from the seminal vesicle, is also known to be required for sperm to escape attack by the female immune system in the uterus (17). In addition, seminal plasma contains various biochemical components that support sperm metabolism, including lipids and organic acids. Differences in the relative size of the accessory reproductive glands reflect differences in the seminal plasma composition. Accordingly, there are considerable differences in the properties of seminal plasma among animal species. Specifically, Rosecrans et al. (18) found significant differences in the levels of most biochemical components in seminal plasma compared to those in blood, with specific factors like fructose being rarely detected in blood but abundant in seminal plasma. Our previous studies (19–21) documented that seminal plasma also contains creatine and fatty acids taken up by sperm within minutes. These substances play crucial roles in regulating sperm metabolic pathways and enhancing sperm motility. These observations indicate that seminal plasma not only alters the female immune system to protect sperm but also directly influences sperm metabolic pathways, facilitating the transformation of sperm capable of in vivo fertilization.

Removed in the revised version: Rosecrans et al. (18) found significant differences in the levels of most biochemical components in seminal plasma compared to blood, with specific factors like fructose being rarely detected in blood but abundant in seminal plasma. Our previous studies (19–21) documented that seminal plasma also contains creatine and fatty acids taken up by sperm within minutes. These substances play crucial roles in regulating sperm metabolic pathways and enhancing sperm motility. These observations indicate that seminal plasma not only alters the female immune system to protect sperm but also directly influences sperm metabolic pathways, facilitating the transformation of sperm capable of in vivo fertilization.

Earlier version: Some specific factors that male accessory glands contribute to seminal plasma and impact male fertility have been elucidated. For example, surgical removal of seminal vesicles in male mice and rats was associated with infertility (17, 22, 23). The observation that fructose (24) and citric acid (25) concentrations in seminal plasma of control mice and rats are higher than in castrated animals indicates that specific functions of the accessory glands might be affected by testicular-derived testosterone that is known to activate the intracellular androgen receptor (AR; NR3C4), required for gene regulation of transcription. AR-deficient mice have been found to be infertile (26), and long-term administration of androgen receptor antagonists, like flutamide, abolished male fertility (27, 28). AR-deficient mice and transcriptome analyses have identified specific targets of AR (29, 30). In both mice and human, androgen levels decrease with aging, and their concentrations have been reported to cause a decrease in male fertility, especially ejaculated sperm motility. Therefore, the reduction of androgen concentrations may affect the metabolic mechanisms of the accessory reproductive glands and alter the composition of seminal plasma with aging. However, the metabolic mechanisms that impact sperm functions are governed by a complex of enzymatic reactions. Therefore, the intricate interplay between testosterone and seminal plasma synthesis in the accessory reproductive glands requires further investigation.

Revised version: Several specific factors produced by the male accessory glands that contribute to seminal plasma and impact male fertility have been elucidated. For example, surgical removal of seminal vesicles in male mice and rats was associated with infertility (17, 22, 23). The observations that fructose (24) and citrate(25) concentrations in seminal plasma of control mice and rats are higher than in castrated animals suggest that the specific metabolism of the accessory glands might be affected by testosterone derived from the testes, which activate intracellular androgen receptors (AR; NR3C4) required for gene regulation of transcription. AR-deficient mice have been found to be infertile (26), and long-term administration of androgen receptor antagonists, like flutamide, abolished male fertility (27, 28). AR-deficient mice and transcriptome analyses have identified specific targets of AR (29, 30). In both mice and humans, androgen levels decrease with aging, and their concentrations have been reported to cause a decrease in male fertility, especially ejaculated sperm motility. In addition to impairments in spermatogenesis, a decline in androgen levels may also disrupt the metabolic functions of the accessory glands and alter the composition of seminal plasma. Therefore, the intricate interplay between testosterone and seminal plasma synthesis in the accessory glands requires further investigation.

Earlier version: In this study, we first performed various bioassays to identify which accessory reproductive glands are essential for sperm fertility. Subsequently, comparative in vivo and in vitro analyses were performed using a hypoandrogenic model. In addition, biochemical and molecular techniques such as RNA-seq, flux analysis, mass spectrometry, and knockdown experiments were employed. These analyses revealed that testosterone promotes the synthesis of oleic acid in seminal vesicle epithelial cells and its secretion into seminal plasma, and the oleic acid ensures the linear motility and fertilization ability of sperm.

Revised version: In this study, various bioassays were done to identify which accessory glands are essential for sperm fertility. Subsequently, comparative in vivo and in vitro analyses were performed using a hypoandrogenic model. In addition, biochemical and molecular techniques such as RNA-seq, flux analysis, mass spectrometry, and knockdown experiments were employed. These analyses revealed that testosterone promotes the synthesis of oleic acid in seminal vesicle epithelial cells and its secretion into seminal plasma.

Results

The effect of androgen-dependent changes in mouse seminal vesicle secretions on the linear motility of sperm

Earlier version: A direct comparison was conducted to clarify the effects of secretions from the prostate and seminal vesicles on the motility parameters of epididymal sperm. Since it is difficult to collect the ejaculated semen of mice, a pseudo-seminal plasma was prepared by preparing fluids from the seminal vesicle secretion and homogenized prostate tissue (Fig.1A, Supplemental Fig.1A). When seminal vesicle secretions were added to the HTF medium, the sperm motility, progressive motile, the straight line velocity (VSL), and the percentage of linearity (LIN) values increased significantly compared to when only prostate extract was used. However, the curvi-linear velocity (VCL) did not change (Fig.1B-F). Meanwhile, the secretions from the prostate and seminal vesicles did not affect the viscosity of the medium (Fig.1G). These results indicate that sperm exposure to seminal plasma factors significantly improves sperm linear motility (high VSL and LIN).

Revised version: A direct comparison was conducted to clarify the effects of secretions from the prostate and seminal vesicles on the motility parameters of epididymal sperm. Since it is difficult to collect the ejaculated semen of mice, a pseudo-seminal plasma was prepared using seminal vesicle secretions and homogenized prostate tissue (Fig.1A, Supplemental Fig.1A). When seminal vesicle secretions were added to the HTF medium, the sperm motility, progressive motile, the straight line velocity (VSL), and the percentage of linearity (LIN) values increased significantly compared to when only prostate extract was used. However, the curvi-linear velocity (VCL) did not change (Fig.1B-F). Meanwhile, the secretions from the prostate and seminal vesicles did not affect the viscosity of the medium (Fig.1G). These results indicate that sperm exposure to seminal plasma factors significantly improved sperm linear motility (high VSL and LIN).

Effects of seminal vesicle fluid and prostate fluids on sperm motility

(A) The effects of secretions from the prostate and seminal vesicles on sperm motility were directly compared. (B-F) The effects of secretions from the prostate (Pros) or seminal vesicles (SV) on the motility parameters of epididymal sperm were tested: (B) Motile, (C) Progressive motile, (D) Curvi-Linear Velocity; VCL, (E) Straight-Line Velocity; VSL and (F) Linearity; LIN is the ratio of VSL to VCL of sperm that was incubated with the mixture of seminal vesicle secretions or prostate extracts. (G) Viscosity of a solution containing prostate or seminal vesicle secretion with a protein concentration of 10 mg/mL. (H) Experimental design to evaluate the effect of seminal vesicle secretions collected from male mice treated with or without flutamide (50 mg/kg subcutaneously every day for 7 days) on sperm. Note that the donor mice for sperm and seminal vesicle secretions were different. (I-L) Performed bioassay using seminal vesicle fluid after treatment with vehicle (Ctrl) or flutamide: (I) Motile, (J) Progressive motile, (K) LIN, (L) VSL. (M) The high mitochondrial membrane potential (hMMP) was checked by the JC-1 kit. Antimycin (Anti) was used as the negative control for hMMP. Data are mean ± SEM. n=3-9 independent replicates. The viscosity measurement data were from n=3 repeated experiments using pooled prostate or seminal vesicle extracts from at least 20 mice. Percentage data were subjected to arcsine transformation before statistical analysis. (B-G, I-L) Significance was tested in comparison using Student’s t-test. (M) Since the results of the Bartlett test were significant, a two-way ANOVA using a generalized linear model was performed. Both the flutamide administration and antimycin addition were significant. The interaction was not significant. Games-Howell was performed as a post hoc test. Different letters represent significantly different groups.

Earlier version: Androgens regulate the functions of seminal vesicles. Therefore, to elucidate whether the seminal vesicle factors, which enhance sperm linear motility, are synthesized in an androgen- dependent manner, bioassays were performed using seminal vesicle secretions collected from control mice or mice treated daily for 7 consecutive days with flutamide, an androgen receptor antagonist (Fig. 1H, Supplemental Fig.1B). The results showed that Motile, Progressive motile, LIN, and VSL were significantly reduced in the sperm cultured in seminal vesicle secretions from flutamide-treated (Flutamide) compared to the sperm cultured in seminal vesicle secretions from controls (Ctrl) (Fig.1I-L). The reduction of VSL, Motile, and Progressive motile were also significantly lower in the sperm treated with seminal vesicle secretions from over 12 months old mice that showed low circular testosterone levels as compared with the sperm treated with seminal vesicle secretions from young adult mice (3 months old mice) (Supplemental Fig.2A,B). Furthermore, JC-1 staining to measure mitochondrial membrane potential, which is important for sperm linear motility, showed a significantly lower percentage of high mitochondrial membrane potential in sperm treated with the secretions of seminal vesicles from flutamide-treated mice and over 12 months old mice (Fig.1M, Supplemental Fig.1C,D, Supplemental Fig.2C). Since flutamide treatment and aging also affect spermatogenesis and sperm motility, it should be noted that epididymal sperm donors differ from seminal vesicle donors (Fig. 1H). The baseline information for the epididymal sperm used in the experiment showed that the sperm were normal (Supplemental Fig.1A). These results indicated that androgen-dependent changes in seminal vesicle functions alter the levels/activity of seminal vesicle secretion factors, which regulate sperm linear motility. The changes in seminal vesicle functions were also induced by aging.

Revised version: Androgens regulate the functions of seminal vesicles. Therefore, to elucidate whether the seminal vesicle factors, which enhance sperm linear motility, are synthesized in an androgen-dependent manner, bioassays were performed using seminal vesicle secretions collected from control mice or mice treated daily for 7 consecutive days with flutamide, an androgen receptor antagonist (Fig. 1H, Supplemental Fig.1B). The results showed that Motile, Progressive motile, LIN, and VSL were significantly reduced in the sperm cultured in seminal vesicle secretions from flutamide-treated (Flutamide) compared to the sperm cultured in seminal vesicle secretions from controls (Ctrl) (Fig.1I-L). The reduction of VSL, Motile, and Progressive motile were also significantly lower in the sperm treated with seminal vesicle secretions from over 12 months old mice that showed low circulating testosterone levels as compared with the sperm treated with seminal vesicle secretions from young adult mice (3 months old mice) (Supplemental Fig.2A,B). Furthermore, JC-1 staining to measure mitochondrial membrane potential, which is important for sperm linear motility, showed a significantly lower percentage of high mitochondrial membrane potential in sperm treated with the secretions of seminal vesicles from flutamide-treated mice and over 12 months old mice (Fig.1M, Supplemental Fig.1C,D, Supplemental Fig.2C). Since flutamide treatment and aging also affect spermatogenesis and sperm motility, it should be noted that epididymal sperm donors differ from seminal vesicle donors (Fig. 1H). Due to the strong effect of the decapacitation factor, it was found that seminal vesicle fluid, regardless of the presence or absence of androgens, reduces the fertilization rate in in vitro fertilization (Supplemental Figure 1E). The baseline information for the epididymal sperm used in the experiment showed that the sperm were normal (Supplemental Fig.1A). These results indicated that androgen-dependent changes in seminal vesicle functions alter the levels/activity of seminal vesicle secretion factors, which regulate sperm linear motility. The changes in seminal vesicle functions were also induced by aging.

The testosterone-androgen receptor pathway inhibits the cell proliferation of seminal vesicle epithelial cells.

(A) The localization of the androgen receptor (AR; NR3C4) in seminal vesicle of control (Ctrl) mice and those treated with flutamide (Fult). Scale bar=50 µm from 10× magnification for left panel and 20 µm from 20× magnification for right panel in each group. (B,C) Cell proliferation and apoptosis in seminal vesicle of control (Ctrl) mice and those treated with flutamide (Fult): (B) Representative images of Ctrl and Flut staining for Ki67 (top) and TUNEL (bottom). Scale bar=20 µm. (C) Percentage of Ki67 positive cells (top) and TUNEL positive (bottom) in Ctrl and Flut-treated seminal vesicle sections. (D) Experimental design to evaluate the effects of testosterone on the proliferation of seminal vesicle epithelial cells in vitro. The epithelial cells were culutured with or without testosterone for 8 days. (E) Growth curves of seminal vesicle epithelial cells that cultured with 0 (c), 1, 10 and 100 ng/mL of testosterone. (F,G) Cell cycle status was determined by flow cytometry. Data are mean ± SEM. n=3-5 mice or independent replicates. Each replicate experiments with 3-6 wells containing pooled cells from 3-5 mice. Percentage data were subjected to arcsine transformation before statistical analysis. Significance was tested in comparison to Ctrl using Student’s t-test. The cell number data at d8 was compared with the control using Dunnett’s test.

Testosterone-androgen receptor activity suppresses the proliferation of seminal vesicle epithelial cells

Earlier version: To understand the difference between seminal vesicles that secrete factors for enhancing sperm linear motility and those that are impaired in this function, we performed histological analysis of seminal vesicles recovered from mice treated and non-treated with flutamide. In control mice, the seminal vesicle epithelial cells were in a single layer, and androgen receptors were localized in the nuclei of seminal vesicle epithelial cells (Fig.2A). However, in mice treated with continuous administration of flutamide, seminal vesicle epithelial cells were multilayered, and the strong accumulation of androgen receptors in nucleus was not observed (Fig.2A). To investigate the cause of this multilayering of cells, we performed Ki67 staining, a marker of cell proliferation, and apoptosis detection by TUNEL staining (Fig.2B). The results showed that the percentage of Ki67- positive cells increased significantly after treatment with flutamide, but the percentage of TUNEL- positive cells did not change, suggesting that the testosterone-androgen receptor pathway is responsible for the inhibition of proliferation of seminal vesicle epithelial cells (Fig.2C). These changes of abnormal morphology and high proliferation activity were also observed in seminal vesicle of more than 12 months old mice that showed the low level of testosterone in serum (Supplemental Figure 2D-F). Therefore, seminal vesicle epithelial cells were collected and cultured in primary culture, as shown in Fig. 2D. The changes in cell number were counted, and the cell cycle was analyzed by flow cytometry. When cultured without testosterone (c), the number of seminal vesicle epithelial cells increased day-dependent manner. However, the addition of testosterone significantly decreased the number of seminal vesicle epithelial cells during the 8-day culture period in a concentration-dependent manner (Fig.2E). Flow cytometric analysis of the cell cycle on day 7 showed that the G0/G1 phase was significantly increased. The G2/M1 phase was significantly decreased by the addition of testosterone (Fig. 2F,G). These results indicate that the testosterone-androgen receptor pathway inhibits seminal vesicle epithelial cells in cell division.

Revised version: To understand the difference between seminal vesicles that secrete factors for enhancing sperm linear motility and those that are impaired in this function, we performed histological analysis of seminal vesicles recovered from mice treated and non-treated with flutamide. In control mice, the seminal vesicle epithelial cells were in a single layer, and androgen receptors were localized in the nuclei of seminal vesicle epithelial cells (Fig.2A). However, in mice treated with continuous administration of flutamide, seminal vesicle epithelial cells were multilayered, and the strong accumulation of androgen receptors in nucleus was not observed (Fig.2A). To investigate the cause of this multilayering of cells, we performed Ki67 staining, a marker of cell proliferation, and apoptosis detection by TUNEL staining (Fig.2B). The results showed that the percentage of Ki67-positive cells increased significantly after treatment with flutamide, but the percentage of TUNEL-positive cells did not change, suggesting that the testosterone-androgen receptor pathway is responsible for the inhibition of proliferation of seminal vesicle epithelial cells (Fig.2C). These changes of abnormal morphology and high proliferation activity were also observed in seminal vesicle of more than 12 months old mice that showed the low level of testosterone in serum (Supplemental Figure 2D-F). Therefore, seminal vesicle epithelial cells were collected and cultured in primary culture, as shown in Fig. 2D. The changes in cell number were counted, and the cell cycle was analyzed by flow cytometry. When cultured without testosterone (c), the number of seminal vesicle epithelial cells increased time-dependent manner. However, the addition of testosterone significantly decreased the number of seminal vesicle epithelial cells during the 8-day culture period in a concentration-dependent manner (Fig.2E). Flow cytometry on day 7 revealed a significant increase in the proportion of cells in the G0/G1 phase and a corresponding decrease in the G2/M phase. The G2/M phase was significantly decreased by the addition of testosterone (Fig. 2F,G). These results indicate that the testosterone-androgen receptor pathway inhibits seminal vesicle epithelial cells in cell division.

Testosterone alters metabolic gene expression in seminal vesicle epithelial cells

To better understand the roles of testosterone in suppressing the proliferation of seminal vesicle epithelial cells, RNA sequencing was done using the epithelial cells cultured with or without testosterone. DEG (Differential Expressed Genes) analysis revealed that a total of 23,459 genes were identified, including 997 upregulated genes and 3463 downregulated genes in seminal vesicle epithelial cells by testosterone relative to control (Fig.3A). Genes largely upregulated by testosterone included those already reported in seminal vesicles (17, 22, 31), such as the Seminal Vesicle Secretion (Svs) family (Svs1∼6), Pate4, Spinkl, Ceacam10, and Sva (Fig.3B,C). Therefore, the primary culture of seminal vesicle epithelial cells used in this experiment was considered to maintain the characteristics of seminal vesicle epithelial cells in vivo. Many genes were downregulated by testosterone, and those with a high rate of variation included collagen synthase and Mmp2, which are involved in extracellular matrix formation (Fig.3C). However, KEGG analysis, rather than individual genes, was used to predict functions altered by testosterone. It was found that the expression of genes involved in the metabolic pathway was the most highly affected by testosterone (Fig.3D). Because there is a close relationship between cell proliferation and metabolic activity and between metabolic substrates in seminal plasma and sperm motility, we hypothesized that testosterone induced changes in seminal vesicle epithelial cell metabolism alter the components of seminal plasma to mediate the enhancement of sperm motility.

Testosterone changes the expression of genes involved in metabolic pathway in seminal vesicle epithelial cells.

(A) Volcano plot of differentially expressed genes. RNA sequencing was performed using RNA extracted from the seminal vesicle epithelial cells cultured with or without 100 ng/mL testosterone. Genes with a significant expression change are highlighted as red dots. It found 4460 significant genes for a cutoff of P<0.05. (B) MA plot of differential expression results. The seminal vesicle specific genes are highlighted as red dots rather than significant genes. (C) Top 10 genes upregulated or downregulated by testosterone, respectively, for altered gene expression. (D) Conducted KEGG analysis to identify differences in pathway enrichment by testosterone, identifying the most variability in metabolic pathway genes. It shows the number of differential expressed genes annotated to Gene Ontology shows (x axis). p-value which measure the statistical significance of a possible functional enrichment for each term.

Earlier version: Testosterone shifts the metabolic activity of seminal vesicle epithelial cells

Revised version: Testosterone shifts the metabolic flux of seminal vesicle epithelial cells

Earlier version: In the presence of testosterone, the addition of glucose significantly activated the extracellular oxidation rate (ECAR) normalized by the total number of live cells compared to that in epithelial cells cultured without testosterone (Fig.4A-C). However, there was no difference in the Oligomycin-sensitive ECAR, indicating that testosterone may increase glucose metabolism but does not enhance the expression of a group of enzymes involved in the glycolytic pathway. To determine the metabolic fate of glycolytic end products, seminal vesicle epithelial cells were cultured with and without testosterone. Subsequently, the concentrations of pyruvate and lactate in the culture supernatant normalized by the total number of live cells were measured. The addition of testosterone significantly decreased the pyruvate concentration; however, the lactate concentration did not change significantly (Fig.4D,E). Therefore, we analyzed mitochondrial metabolism using a flux analyzer, predicting that more glucose is metabolized, pyruvate is synthesized in response to testosterone, and is further metabolized in the mitochondria. Unexpectedly, when measuring mitochondrial respiratory metabolism normalized by the total number of live cells with a Flux analyzer, the basal respiratory rate (the electron transport chain activity) was significantly decreased by testosterone (Fig.4F-J). Oligomycin-sensitive respiration was also significantly suppressed by exposure to testosterone. Furthermore, the FCCP uncoupled respiration and spare respiration capacity were also significantly decreased in response to testosterone treatment. These results suggest that testosterone does not reduce transient enzyme activity in mitochondria but rather weakens the metabolic pathway of the mitochondrial TCA cycle and/or the electron transport chain due to the changes in gene expression patterns in seminal vesicle epithelial cells.

Revised version: In the presence of testosterone, the addition of glucose significantly activated the extracellular acidification rate (ECAR) normalized by the total number of live cells compared to that in epithelial cells cultured without testosterone (Fig.4A-C). These data indicate that testosterone enhances glucose utilization. This leads to the interpretation that testosterone increases the flow of glycolysis by increasing glucose uptake and alters metabolic flux distribution. To determine the metabolic fate of glycolytic end products, seminal vesicle epithelial cells were cultured with and without testosterone. Subsequently, the concentrations of pyruvate and lactate in the culture supernatant normalized by the total number of live cells were measured. The addition of testosterone significantly decreased the pyruvate concentration; however, the lactate concentration did not change significantly (Fig.4D,E). We therefore examined cellular energy metabolism with a flux analyzer, anticipating that testosterone would elevate glycolytic flux, thereby producing more pyruvate from phosphoenolpyruvate. Because extracellular pyruvate levels simultaneously declined, we inferred that the cells had an increased pyruvate demand and, at that time, hypothesized that the excess pyruvate would enter the mitochondria to support enhanced oxidative metabolism. Unexpectedly, when measuring mitochondrial respiratory metabolism normalized by the total number of live cells with the flux analyzer, the basal respiratory rate (the electron transport chain activity) was significantly decreased by testosterone (Fig.4F-J). Oligomycin-sensitive respiration was also significantly suppressed by exposure to testosterone. Furthermore, the FCCP uncoupled respiration and spare respiration capacity were also significantly decreased in response to testosterone treatment. Overall, these data indicate that testosterone promotes the redistribution of metabolic flux. In other words, testosterone increased glycolysis in seminal vesicle epithelial cells while decreasing mitochondrial respiration. To determine whether these changes were accompanied by changes in gene expression of specific metabolic-related enzymes, we analyzed gene expression levels.

Testosterone regulates glucose metabolism and mitochondrial ATP production in seminal vesicle epithelial cells.

(A-C) Extracellular acidification rate (ECAR) analysis by an extracellular flux analyzer in seminal vesicle epithelial cells cultured with 100 ng/mL testosterone (Testo) or vehicle (Ctrl) for 7 days: (A) ECAR kinetics of seminal vesicle epithelial cells using an extracellular flux analyzer. (B) Glucose induced ECAR and (C) Oligomycin-sensitive ECAR. (D,E) The concentrations of pyruvate and lactate in the culture supernatant. These measurements were normalized based on cell count and viability. (D) Pyruvate concentration in the medium where seminal vesicle epithelial cells were cultured with or without testosterone for 24 hours. (E) Lactate concentration. (F-J) Mitochondrial respiration measurement by an extracellular flux analyzer: (F) Oxygen consumption rate (OCR) kinetics of seminal vesicle epithelial cells with or without 100 ng/mL testosterone. (G) Basal OCR, (H) FCCP uncoupled respiration, (I) Spare respiratory capacity and (J) Oligomycin-sensitive respiration. Data are mean ± SEM of n = 3 replicate experiments with 3-6 wells containing pooled cells from 3-5 mice or medium. Student’s t-test was used to compare Ctrl and Testo. The cells were normalized to 10,000 viable cells/well immediately before analysis. The viability of the cells before experiments was 86-93%.

Testosterone causes changes in the expression of genes encoding enzymes involved in glucose catabolism and anabolism

Earlier version: The expression of genes involved in glucose catabolism and anabolism in the in vitro culture system of seminal vesicle epithelial cells was analyzed by qPCR (Fig.5). The gene expression of glucose transporter (Glut) and enzymes involved in the glycolytic pathway was not significantly affected by testosterone, consistent with the results of metabolic flux analysis, which showed that the glycolytic pathway was not enhanced despite the activation of the glycolysis. Furthermore, testosterone did not significantly affect the expression of genes encoding enzymes involved in the TCA cycle in mitochondria. However, the expression of enzymes involved in the electron transport chain was significantly decreased. Interestingly, the expression of Acly, which encodes an enzyme that reconverts citric acid formed in the TCA cycle to Acetyl CoA and induces cholesterol and fatty acid synthesis, was significantly increased by testosterone. Although Hmgcr, which encodes the rate-limiting enzyme involved in cholesterol synthesis, was unchanged, the expression of Acc, which encodes the rate-limiting enzyme for fatty acid synthesis, was significantly upregulated by testosterone. These characteristic testosterone-induced changes in gene expression observed in cultured cells were also detected by ACLY immunostaining of in vivo seminal vesicle samples (Supplemental Figure 3). Specifically, the expression of genes encoding a group of enzymes of the electron transport chain was significantly upregulated by continuous flutamide administration in vivo, whereas the gene expression of Acly and Acc was significantly decreased by flutamide administration. Immunostaining results showed that ACLY was detected specifically in seminal vesicle epithelial cells, and their staining appeared to be reduced by flutamide treatment. The lower level of ACLY was also observed in seminal vesicle epithelial cells of over 12 months old mice (Supplemental Figure 3B).

Revised version: The expression of genes involved in glucose catabolism and anabolism in the in vitro culture system of seminal vesicle epithelial cells was analyzed by qPCR (Fig.5). The gene expression of glucose transporter (Glut) and enzymes involved in the glycolytic pathway was not significantly affected by testosterone, consistent with the results of metabolic flux analysis, which showed that the glycolytic pathway was not enhanced despite the activation of the glycolysis. Furthermore, testosterone did not significantly affect the expression of genes encoding enzymes involved in the TCA cycle in mitochondria. However, the expression of enzymes involved in the electron transport chain was significantly decreased. Interestingly, testosterone significantly increased the expression of Acly, which encodes a cytoplasmic enzyme that converts citrate transported from the TCA cycle into acetyl-CoA, a substrate that is essential for fatty acid synthesis. Although Hmgcr, which encodes the rate-limiting enzyme involved in cholesterol synthesis, was unchanged, the expression of Acc, which encodes the rate-limiting enzyme for fatty acid synthesis, was significantly upregulated by testosterone. These characteristic testosterone-induced changes in gene expression observed in cultured cells were also detected by ACLY immunostaining of in vivo seminal vesicle samples (Supplemental Figure 3). Specifically, the expression of genes encoding a group of enzymes of the electron transport chain (mtNd6) was significantly upregulated by continuous flutamide administration in vivo, whereas the gene expression of Acly and Acc was significantly decreased by flutamide administration (Supplemental Figure 3B). Immunostaining results showed that ACLY was detected specifically in seminal vesicle epithelial cells, and their staining appeared to be reduced by flutamide treatment. The lower level of ACLY was also observed in seminal vesicle epithelial cells of over 12 months old mice (Supplemental Figure 3C).

Effect of testosterone on gene expression of enzymes involved in the glucose metabolic pathway in seminal vesicle epithelial cells.

To elucidate the effects of testosterone on gene expression of enzymes involved in glucose catabolism and anabolism in seminal vesicle epithelial cells. ctrl: 7 days of culture with vehicle. Testo: 7 days of culture with 100 ng/mL testosterone. Data are mean ± SEM of n = 3 replicate experiments with 3-6 wells containing pooled cells from 3-5 mice. Student’s t-test was used to compare Ctrl and Testo.

Mechanism of testosterone-induced enhancement of glucose uptake and, thereby, fatty acid synthesis

Since it is known that testosterone regulates the intracellular localization of GLUT4, we focused on the intracellular localization of GLUT4 and its glucose uptake capacity (32, 33). The effect of testosterone on glucose uptake in seminal vesicle epithelial cells was analyzed using fluorescence-labeled glucose. As expected from the results of flux metabolic analyses, testosterone treatment significantly increased glucose uptake (Fig.6A). Indinavir, a functional inhibitor of GLUT4, significantly suppressed glucose uptake and significantly decreased fatty acid synthesis (the amount of fatty acids accumulated in the cell), which was increased by testosterone (Fig.6A,B).

Testosterone enhances glucose uptake and fatty acid synthesis via GLUT4 trans-localization in seminal vesicle epithelial cells.

(A) The glucose uptake ability of seminal vesicle epithelial cells was detected using fluorescence-tagged glucose. The cells were cultured with 100 ng/mL testosterone and/or indinavir (100 µM) for 7 days and then further treated with glucose uptake probe-green for 15 min, and intracellular fluorescence was measured by flow cytometry (n=4). (B) Lipid accumulation in the medium where seminal vesicle epithelial cells were incubated for 24 hours (n=6). (C) Immunostaining of GLUT4 in flutamide-treated or vehicle mice derived seminal vesicle. Scale bar is 50 µm. (D) Western blot images of GLUT4 and α/β-tubulin in three sets of seminal vesicle epithelial cells cultured with 100 ng/mL testosterone (Testo) or in vehicle (Ctrl) for 7 days. Quantitative analysis of GLUT4 relative to α-tubulin obtained from Western blot. (E) Fatty acid composition in the medium was analyzed using gas chromatography. (F) Quantification of Elovl family gene expression by RT-qPCR. (G) Pathway of glucose assimilation to oleic acid. Data are mean ± SEM. n=3-6 independent replicates. Each replicate experiments with 3-6 wells containing pooled cells from 3-5 mice or medium. (A,B) The glucose uptake and lipid measurements were normalized based on cell count and viability. A two-way ANOVA was performed. Both the Indinavir and testosterone treatment were significant. The interaction was significant. Tukey’s Honest Significant Difference test was performed as a post hoc test. Different letters represent significantly different groups. Student’s t-test was used for comparison between the two groups.

Earlier version: Furthermore, immunostaining analysis of the localization of GLUT4 in seminal vesicles revealed that GLUT4 was translocalized outside the nucleus of seminal vesicle epithelial cells under the presence of testosterone (Fig.6C). These data suggest that GLUT4 was partially localized at the plasma membrane and contributed to glucose uptake. In contrast, in seminal vesicles from mice treated with flutamide continuously, no strong signal was detected at the presumed location of GLUT4 near the plasma membrane of seminal vesicle epithelial cells, even though the protein expression level was not affected (Fig.6C,D). The same characteristics were observed in aged mice.

Revised version: Furthermore, immunostaining analysis of the localization of GLUT4 in seminal vesicles revealed that GLUT4 was translocated outside the nucleus of seminal vesicle epithelial cells under the presence of testosterone (Fig.6C). These data suggest that GLUT4 was partially localized at the plasma membrane and contributed to glucose uptake. In contrast, in seminal vesicles from mice continuously treated with flutamide, no strong signal was detected at the presumed location of GLUT4 near the plasma membrane of seminal vesicle epithelial cells, even though the protein expression level was not affected (Fig.6C,D). The same characteristics were observed in aged mice.

Earlier version: Qualitative and quantitative analysis of the fatty acids synthesis was performed by GC-MS. In seminal vesicles, C16:0, C18:0, and C18:1 were detected, in which C18:1 was decreased to an undetectable level in seminal vesicles of mice treated with the flutamide (Supplemental Figure 4A). In the cultured seminal vesicle epithelial cells, C18:1, oleic acid, was significantly increased by the addition of testosterone as observed in the seminal vesicle (Fig.6E). The expression of Elovl6 encoding ELOVL6, which desaturates C18:0 saturated fatty acids to C18:1, was significantly up-regulated by testosterone (Fig.6F,G). The testosterone-induced change was also detected in in vivo samples (seminal vesicle), where the expression of Evovl6 was decreased by the flutamide treatment (Supplemental Figure 4B).

Revised version: Qualitative and quantitative analysis of the fatty acids synthesis was performed by GC-MS. In seminal vesicles, C16:0, C18:0, and C18:1 were detected, in which C18:1 was decreased to an undetectable level in seminal vesicles of mice treated with the flutamide (Supplemental Figure 4A). In the cultured seminal vesicle epithelial cells, C18:1, oleic acid, was significantly increased by the addition of testosterone as observed in the seminal vesicle (Fig.6E). The expression of Elovl6 encoding ELOVL6, which desaturates C18:0 saturated fatty acids to C18:1, was significantly up-regulated by testosterone (Fig.6F,G). The testosterone-induced change was also detected in in vivo samples (seminal vesicle), where the expression of Elovl6 was decreased by the flutamide treatment (Supplemental Figure 4B).

Added in the revised version: In order to examine whether these lipids are the factors that induces sperm linear motility, we examined whether or not the addition of lipid mixture changed in sperm motility pattern (Supplemental Figure 5). The lipid mixture did not affect sperm motility at any concentration compared to the control (LM 0%). However, VSL was significantly increased in sperm with 1% lipid mixture added.

Removed in the revised version: Oleic acid (18:1) is taken up by sperm and promotes their linear motility

Removed in the revised version: We observed that seminal vesicles secrete testosterone-dependent factors that induce sperm to move linear and that seminal vesicle epithelial cells synthesize and secrete oleic acid in a testosterone-dependent manner. In order to examine whether oleic acid is a factor that induces sperm linear motility, we examined whether or not oleic acid is taken up by sperm and then changes in sperm motility pattern. When epididymal sperm were incubated with fluorescently labeled oleic acid, the fluorescent signal was selectively detected in the midpiece of the sperm after quenching of extracellular fluorescence (Fig.7A). The fluorescence intensity of sperm detected by flow cytometry significantly increased in a dose-dependent manner (Fig.7B, C). Oleic acid did not affect sperm and progressive motile at any concentration compared to the control (Fig.7D,E). In oleic acid-treated sperm, VSL was significantly increased in a dose- dependent manner in the 0 to 50 nM (Fig.7F). LIN values were highest when 10 nM oleic acid was added (Fig.7G). A flux analyzer was used to monitor real-time changes in mitochondrial β- oxidation normalized by the number of sperm to quantify metabolic changes in mitochondrial activity. The injection of 10 nM oleic acid induced a significant increase in OCR. Etomoxir, a β- oxidation inhibitor, significantly decreased the OCR that was induced by oleic acid (Fig.7H, Supplemental Figure 5A). In situ ATP detection by BioTracker ATP-Red revealed that the percentage of sperm with high ATP production was significantly increased under the presence of oleic acid (Fig.7I). To investigate the effect of oleic acid on the linear motility of sperm and its effect on the fertilization ability of sperm, in vitro fertilization (IVF) was performed. When the sperm number was sufficient, there was no change in the cleavage rate after in vitro fertilization, but when the sperm count was reduced to less than one-tenth of the conventional level, the cleavage rate increased compared to the control (Fig.7J,K). Due to the strong effect of the decapacitation factor, we found that seminal vesicle fluid reduces the fertilization rate in IVF (Supplemental Figure 5B). This results from inhibiting tyrosine phosphorylation of sperm and acrosome reaction by seminal vesicle secretions. However, oleic acid did not have this inhibitory effect (Supplemental Figure 5C-F).

Removed in the revised version: Oleic acid is incorporated to sperm, which enhances linear motility and mitochondrial activity.

Removed in the revised version: (A) Epididymal sperm were incubated with fluorescently labeled oleic acid (OA) and the fluorescence intensity after quenching was observed by a fluorescence microscope. (B) The fluorescence intensity of fluorescence-conjugated OA in sperm was detected by flow cytometry. (C) Average fluorescence intensity of sperm at different concentrations of fluorescence-conjugated OA. (D-G) OA improved sperm motility parameters. Sperm collected from mouse epididymis was incubated with OA for 60 min: (D) Motile, (E) Progressive motile, (F) Straight-Line Velocity; VSL and (G) Linearity; LIN (H) Increase in the Oxygen consumption rate (OCR) of cultured sperm was measured using flux analyzer after treatment with 10 nM OA or 10 nM OA + 40 nM Etomoxir (Eto). Ctrl: the sperm cultured in HTF medium. Each well was seeded with 180 μL of NaHCO3-Free HTF medium containing 3,000,000 sperm. (I) The percentage of ATP-Red positive sperm after 1 hour incubation with or without 10 nM OA, measured using BioTracker ATP-Red staining and flow cytometry. (J,K) The sperm, after 1 hour of incubation with or without 10 nM OA, were used for the IVF test and analyzed for the rate of oocytes reaching cleavage. and blastocyst stage. (J) Cleavage rate. (K) Blastocyst rate. (L) Artificial insemination using sperm treated for 1 hour with “pseudo” seminal plasma from control (Ctrl), flutamide-treated (Flu) mice and supplement of 10 nM OA (Flu+OA) followed by an evaluation of the fertilization rates. Note that the donor mice for sperm and seminal vesicle secretions were different. Data are mean ± SEM. n=3-11 independent replicates. Percentage data were subjected to arcsine transformation before statistical analysis. Differences between groups were assessed by one-way analysis of variance (ANOVA). When ANOVA was significant, differences among values were analyzed by Tukey’s Honest Significant Difference test for multiple comparisons. Different letters represent significantly different groups. Student’s t-test was used for comparison between the two groups.

Removed in the revised version: To examine whether the addition of oleic acid to low quality seminal plasma (collected from flutamide-treated mice) restores in vivo the fertilization potential of epididymal sperm, sperm treated with seminal plasma collected from control or flutamide-treated mice was used for artificial insemination. The in vivo fertilization rate was about 70% when sperm was treated with the seminal vesicle secretions from control mice (SV); however, the fertilization rate was significantly lower in the group that sperm was treated with the seminal vesicle secretions from flutamide-treated mice (Flu) (SV: 73.1 ± 3.0%, Flu: 44.9 ± 4.5%). Although not significant, adding oleic acid to the seminal secretions of flutamide-treated mice (Flu+OA) increased the fertilization rate to 61.3 ± 3.9% (Fig.7L). These results indicated that the physiological role of oleic acid is to increase the probability of fertilization by keeping the linear sperm motility. This increases the likelihood of the sperm passing through the female reproductive tract and environments that are unfavorable to sperm survival.

Earlier version: ACLY expression is upregulated by testosterone and is essential for the metabolic shift of seminal vesicle epithelial cells that mediates sperm linear motility

Revised version: ACLY expression is upregulated by testosterone and is essential for the metabolic shift that is associated with increased linear motility

Earlier version: ACLY is upregulated in seminal epithelial cells as a factor that enables the utilization of citric acid from the TCA cycle for fatty acid synthesis outside the mitochondria. Therefore, we performed siRNA knockdown experiments on ACLY to clarify the impact of ACLY on the cellular metabolic shift. The addition of testosterone significantly increased the amount of ACLY protein in seminal epithelial cells (Fig.8A,B). Viral infection with siRNAs encoding Acly- targeted siRNAs markedly decreased the amount of ACLY (Fig.8C). When ACLY was knocked down using shRNA, the proliferation arrest caused by testosterone was released, and the cells began to proliferate again (Fig.8D). This knockdown of ACLY increased the oxygen consumption of seminal epithelial cells in the control and testosterone-treated group (Fig.8E,F). Thus, ACLY expression is associated with mitochondrial metabolic pathway activity in seminal epithelial cells. Furthermore, the knockdown of ACLY also significantly suppressed testosterone-induced fatty acid synthesis in seminal epithelial cells (Fig.8G). According to the GC-MS quantification results, the amount of fatty acids secreted into the culture medium was also reduced by ACLY knockdown. In particular, the enhancement of oleic acid secretion by testosterone was inhibited by ACLY knockdown (Fig.8H-J). Supernatants from epithelial cells of the testosterone-added group induced motility and linear motility (Highest VSL and LIN; Fig.8K-N). However, the knockdown of ACLY significantly reduced the effects of the culture supernatant on sperm motility and LIN (Fig.8K-N). These results found that cell proliferation stops when a metabolic shift mediated by ACLY occurs, and glucose is reallocated to fatty acid synthesis. Fatty acid synthesis is an important function of the seminal vesicles, and it was found that when testosterone is present, fatty acid synthesis enhancement and cell proliferation stop simultaneously. However, it is not known what effects occur when cell proliferation stops.

Revised version: ACLY is upregulated in seminal vesicle epithelial cells as a factor that enables the utilization of citrate from the TCA cycle for fatty acid synthesis outside the mitochondria. Therefore, we performed shRNA knockdown experiments on ACLY to clarify the impact of ACLY on the cellular metabolic shift. The addition of testosterone significantly increased the amount of ACLY protein in seminal vesicle epithelial cells (Fig.7A,B). Viral infection with shRNAs encoding Acly-targeted shRNAs markedly decreased the amount of ACLY (Fig.7C). When ACLY was knocked down using shRNA, the proliferation arrest caused by testosterone was released, and the cells began to proliferate again (Fig.7D). This knockdown of ACLY increased the oxygen consumption of seminal vesicle epithelial cells in the control and testosterone-treated group (Fig.7E,F). Thus, ACLY expression is associated with mitochondrial metabolic pathway activity in seminal vesicle epithelial cells. Furthermore, the knockdown of ACLY also significantly suppressed testosterone-induced fatty acid synthesis in seminal vesicle epithelial cells (Fig.7G). According to the GC-MS quantification results, the amount of fatty acids secreted into the culture medium was also reduced by ACLY knockdown. In particular, the enhancement of oleic acid secretion by testosterone was inhibited by ACLY knockdown (Fig.7H-J). Supernatants from epithelial cells of the testosterone-added group induced motility and linear motility (Highest VSL and LIN; Fig.7K-N). However, the knockdown of ACLY significantly reduced the effects of the culture supernatant on sperm motility and LIN (Fig.7K-N). These results found that cell proliferation stops when a metabolic shift mediated by ACLY occurs, and glucose is reallocated to fatty acid synthesis. Fatty acid synthesis is an important function of the seminal vesicles, and it was found that when testosterone is present, fatty acid synthesis enhancement and cell proliferation stop simultaneously. However, it is not known what effects occur when cell proliferation stops.

Testosterone-regulated ACLY induces metabolic shifts in seminal vesicle epithelial cells.

Earlier version: (A) Western blot images of ACLY and α/β-tubulin in three sets of seminal vesicle epithelial cells cultured with 100 ng/mL testosterone (Testo) or in vehicle (Ctrl) for 7 days. (B) Quantitative analysis of ACLY relative to α-tubulin obtained from Western blot. (C) siRNA knockdown experiments of ACLY in seminal vesicle epithelial cells. ACLY protein levels in scrambled shRNA or ACLY shRNA-transfected seminal vesicle epithelial cells cultured with or without 100 ng/mL testosterone were determined by Western blot. (D) Number of cells at 7 days after incubation. (E,F) Changes in oxygen consumption of seminal epithelial cells were analyzed by a flux analyzer: (E) Oxygen consumption rate (OCR) kinetics of seminal vesicle epithelial cells transfected with scrambled shRNA or ACLY shRNA. (F) Basal OCR. (G) The impact of ACLY knockdown on testosterone-induced fatty acid synthesis was measured. These measurements were normalized based on cell count and viability. The viability of the cells before experiments was 67-81%. (H-I) Fatty acid composition in the cultured supernatants was analyzed using gas chromatography. (H) Palmitate (C16:0). (I) Stearic acid (C18:0). (J) Oleic acid (C18:1). (K-N) The effect of culture supernatants of testosterone-treated cells on sperm motility was evaluated, especially after ACLY knockdown. (K) Motile. (L) Progressive motile. (M) Straight-Line Velocity; VSL. (N) Linearity; LIN. Data are mean ± SEM. n=3-6 independent replicates. Repeated experiments were performed with cells recovered from 3 wells containing pooled cells from 3 to 5 mice. Student’s t-test was used for comparison between the two groups. Percentage data were subjected to arcsine transformation before statistical analysis. (B) Significance was tested in comparison using Student’s t-test. (D,G-M) A two-way ANOVA was performed. Tukey’s Honest Significant Difference test was performed as a post hoc test. (D) Both the ACLY knockdown and testosterone treatment were significant. The interaction was significant. (G) Both the ACLY knockdown and testosterone treatment were significant. The interaction was not significant. (H,I) Only the ACLY knockdown were significant. The interaction was not significant. (J) Both the ACLY knockdown and testosterone treatment were significant. The interaction was significant. (F,N) Since the results of the Bartlett test were significant, a two-way ANOVA using a generalized linear model was performed. Games-Howell was performed as a post hoc test. (E) Only the ACLY knockdown were significant. The interaction was not significant. (N) Testosterone treatment were significant. The interaction was significant. Different letters represent significantly different groups.

Revised version: (A) Western blot images of ACLY and α/β-tubulin in three sets of seminal vesicle epithelial cells cultured with 100 ng/mL testosterone (Testo) or in vehicle (Ctrl) for 7 days. (B) Quantitative analysis of ACLY relative to α-tubulin obtained from Western blot. (C) shRNA knockdown experiments of ACLY in seminal vesicle epithelial cells. ACLY protein levels in scrambled shRNA or ACLY shRNA-transfected seminal vesicle epithelial cells cultured with or without 100 ng/mL testosterone were determined by Western blot. (D) Number of cells at 7 days after incubation. (E,F) Changes in oxygen consumption of seminal vesicle epithelial cells were analyzed by a flux analyzer: (E) Oxygen consumption rate (OCR) kinetics of seminal vesicle epithelial cells transfected with scrambled shRNA or ACLY shRNA. (F) Basal OCR. (G) The impact of ACLY knockdown on testosterone-induced fatty acid synthesis was measured. These measurements were normalized based on cell count and viability. The viability of the cells before experiments was 67-81%. (H-I) Fatty acid composition in the cultured supernatants was analyzed using gas chromatography. (H) Palmitate (C16:0). (I) Stearic acid (C18:0). (J) Oleic acid (C18:1). (K-N) The effect of culture supernatants of testosterone-treated cells on sperm motility was evaluated, especially after ACLY knockdown. (K) Motile. (L) Progressive motile. (M) Straight-Line Velocity; VSL. (N) Linearity; LIN. Data are mean ± SEM. n=3-6 independent replicates. Repeated experiments were performed with cells recovered from 3 wells containing pooled cells from 3 to 5 mice. Student’s t-test was used for comparison between the two groups. Percentage data were subjected to arcsine transformation before statistical analysis. (B) Significance was tested in comparison using Student’s t-test. (D,G-M) A two-way ANOVA was performed. Tukey’s Honest Significant Difference test was performed as a post hoc test. (D) Both the ACLY knockdown and testosterone treatment were significant. The interaction was significant. (G) Both the ACLY knockdown and testosterone treatment were significant. The interaction was not significant. (H,I) Only the ACLY knockdown were significant. The interaction was not significant. (J) Both the ACLY knockdown and testosterone treatment were significant. The interaction was significant. (F,N) Since the results of the Bartlett test were significant, a two-way ANOVA using a generalized linear model was performed. Games-Howell was performed as a post hoc test. (E) Only the ACLY knockdown were significant. The interaction was not significant. (N) Testosterone treatment were significant. The interaction was significant. Different letters represent significantly different groups.

Earlier version: This characteristic metabolic mechanism induced by testosterone is also observed in human seminal epithelial cells

Revised version: This characteristic metabolic mechanism induced by testosterone is also observed in human seminal vesicle epithelial cells

Earlier version: To determine whether these findings in mice could be applied to humans, the commercially available human seminal vesicle epithelial cells were used for following studies (Fig.9A). The growth curves showed that cell proliferation was inhibited by 100 ng/mL of testosterone compared to the control; a pattern similar to that observed in murine seminal vesicle epithelial cells (Fig.9B). The characteristics of metabolic activity of human seminal epithelial cells were also similar to those of mice, with testosterone significantly enhancing the ECAR response (Fig.9C-E), but significantly reducing the capacity for oxygen consumption in mitochondria (Fig.9F-H). Moreover, fluorescent-glucose uptake was increased by testosterone and decreased by the inhibitor of GLUT4 function (Fig.9I). Additionally, while testosterone also significantly increased fatty acid synthesis, this increase was completely abolished by the GLUT4 inhibitor (Fig.9J). According to the quantitative results obtained using GC-MS, the amount of fatty acids secreted into the culture medium also decreased due to the inhibition of GLUT4 by indinavir. In particular, it was shown that oleic acid secretion is completely dependent on glucose via GLUT4 (Fig.9K-M).

Revised version: To determine whether these findings in mice could be applied to humans, the commercially available human seminal vesicle epithelial cells were used for following studies (Fig.8A). The growth curves showed that cell proliferation was inhibited by 100 ng/mL of testosterone compared to the control; a pattern similar to that observed in murine seminal vesicle epithelial cells (Fig.8B). The characteristics of metabolic activity of human seminal vesicle epithelial cells were also similar to those of mice, with testosterone significantly enhancing the ECAR response (Fig.8C-E), but significantly reducing the capacity for oxygen consumption in mitochondria (Fig.8F-H). Moreover, fluorescent-glucose uptake was increased by testosterone and decreased by the inhibitor of GLUT4 function (Fig.8I). Additionally, while testosterone also significantly increased fatty acid synthesis, this increase was completely abolished by the GLUT4 inhibitor (Fig.8J). According to the quantitative results obtained using GC-MS, the amount of fatty acids secreted into the culture medium also decreased due to the inhibition of GLUT4 by indinavir. In particular, it was shown that oleic acid secretion is completely dependent on glucose via GLUT4 (Fig.8K-M).

Earlier version: Testosterone regulates the metabolic activity of human seminal epithelial cells.

Revised version: Testosterone regulates the metabolic activity of human seminal vesicle epithelial cells.

Earlier version: (A) Representative image of human seminal vesicle epithelial cells. (B) Growth curves of seminal vesicle epithelial cells with 100 ng/mL testosterone (Testo) or without (Ctrl). (C-H) Extracellular acidification rate (ECAR) kinetics and Oxygen consumption rate (OCR) kinetics of human seminal vesicle epithelial cells after 6 days of culture with or without 100 ng/mL testosterone using a flux analyzer. (D) Glucose induced ECAR and (E) Oligomycin-sensitive ECAR. (G) Basal OCR, (H) FCCP uncoupled respiration. (I) The glucose uptake ability of seminal vesicle epithelial cells was detected using fluorescence-tagged glucose. (J) Lipid accumulation in the medium where human seminal vesicle epithelial cells were incubated for 24 hours. (H-I) Fatty acid composition in the cultured supernatants was analyzed using gas chromatography. (H) Palmitate (C16:0). (I) Stearic acid (C18:0). (J) Oleic acid (C18:1). Data are mean ± SEM. n=3 independent replicates. Student’s t-test was used for comparison between the two groups. (C-J) These measurements were normalized based on cell count and viability. The viability of the cells before experiments was 88-96%. (I,K,L) Since the results of the Bartlett test were significant, a two-way ANOVA using a generalized linear model was performed. Games-Howell was performed as a post hoc test. Both the Indinavir and testosterone treatment were significant. The interaction was significant. (J) A two-way ANOVA was performed. Tukey’s Honest Significant Difference test was performed as a post hoc test. Both the Indinavir and testosterone treatment were significant. The interaction was significant. Different letters represent significantly different groups.

Revised version: (A) Representative image of human seminal vesicle epithelial cells. (B) Growth curves of seminal vesicle epithelial cells with 100 ng/mL testosterone (Testo) or without (Ctrl). (C-H) Extracellular acidification rate (ECAR) kinetics and Oxygen consumption rate (OCR) kinetics of human seminal vesicle epithelial cells after 6 days of culture with or without 100 ng/mL testosterone using a flux analyzer. (D) Glucose induced ECAR and (E) Oligomycin-sensitive ECAR. (G) Basal OCR, (H) FCCP uncoupled respiration. (I) The glucose uptake ability of seminal vesicle epithelial cells was detected using fluorescence-tagged glucose. (J) Lipid accumulation in the medium where human seminal vesicle epithelial cells were incubated for 24 hours. (K-M) Fatty acid composition in the cultured supernatants was analyzed using gas chromatography. (K) Palmitate (C16:0). (L) Stearic acid (C18:0). (M) Oleic acid (C18:1). Data are mean ± SEM. n=3 independent replicates. Student’s t-test was used for comparison between the two groups. (C-J) These measurements were normalized based on cell count and viability. The viability of the cells before experiments was 88-96%. (I,K,L) Since the results of the Bartlett test were significant, a two-way ANOVA using a generalized linear model was performed. Games-Howell was performed as a post hoc test. Both the Indinavir and testosterone treatment were significant. The interaction was significant. (J) A two-way ANOVA was performed. Tukey’s Honest Significant Difference test was performed as a post hoc test. Both the Indinavir and testosterone treatment were significant. The interaction was significant. Different letters represent significantly different groups.

Discussion

Seminal plasma has a variety of functions, including sperm transport, nutritional support, interactions with the female reproductive tract, and especially the regulation of its relationship with the female immune system. Mammalian seminal plasma is secreted mainly from the seminal vesicle and is rich in cytokines, prostaglandins, and SVS family members (34). For this reason, recent studies have focused on the immunomodulatory function of seminal plasma in the female reproductive tract. On the other hand, fructose and lipids in seminal plasma are the major energy sources for sperm (35). However, although metabolomic analyses of seminal plasma have shown that metabolic substrates are present at higher concentrations in seminal plasma than in blood (18, 24, 25), how these substrates are synthesized and specifically how they alter sperm motility has not been determined. Moreover, though the motility of ejaculated spermatozoa is known to decrease with aging, few reports have revealed age-related changes in the seminal vesicles, which are responsible for seminal plasma synthesis.

In this study, we focused on the reproductive organs that synthesize seminal vesicle secretion as target organs for testosterone and found that factors formed by testosterone-dependent metabolic changes in seminal vesicles activate sperm mitochondria and enhance their linear motility. Testosterone is known to activate intracellular ARs and alter intracellular metabolic pathways in skeletal muscle (36), prostate cells (37) and cardiomyocytes (33). ARs are known to directly regulate the expression of glycolytic and lipid metabolic genes (33, 38, 39). In the present study, the nuclear localization of AR in seminal vesicle epithelial cells (i.e., functioning as a transcription factor) and the activation of the glycolytic system by testosterone were similar to those in the above cell types. In seminal vesicle epithelial cells, an increased trend in the expression of Glut3 and an enhancement of GLUT4 function were observed in a testosterone-dependent manner. Such effects of testosterone on glucose metabolism have also been reported in skeletal muscle, liver, and adipose tissue, where testosterone leads to the phosphorylation of GLUT4 and enhances its translocation to the plasma membrane (40, 41). This is consistent with the results of our metabolic flux analyses, which show that testosterone increases the rate of glycolysis in seminal vesicle epithelial cells by increasing glucose uptake. We further confirmed that GLUT4 is the transporter that mediates this glucose uptake using indinavir, a GLUT4-specific inhibitor.

This study suggests that fatty acid synthesis is an important function of the seminal vesicles, which provide an environment where the ejaculated sperm can carry out active metabolism. The shift to fatty acid synthesis was consistent with the cessation of cell proliferation in the presence of testosterone. This mechanism could prevent the seminal vesicles from excessively consuming the substrates necessary for sperm motility. Manipulation of ACLY expression in mouse myoblasts and embryonic stem cells has been shown to alter the TCA cycle into a “non-normal cycle” (42). One advantage of the non-normal TCA cycle is its ability to retain carbon instead of combusting it and to regenerate the cytosolic NAD+ needed to maintain glycolysis (42, 43). It suggests that ACLY, activated by testosterone, may suppress proliferation by skipping several steps in the mitochondrial TCA cycle and minimizing ATP production associated with aerobic respiration in seminal vesicle epithelial cells.

Earlier version: Testosterone has been reported to cause hypertrophy of cardiac myocytes (44, 45). Cell hypertrophy is partly due to the accumulation of fatty acids in the cytoplasm, which causes significant damage to cellular functions (46, 47). On the other hand, normal cell hypertrophy, especially in myofibers, involves the activation of signaling molecules such as Myc, hypoxia-inducible factor, and Pi3k-Akt-mTor, which regulate metabolic reprogramming in cancer by increasing glucose uptake in the presence of oxygen (48–50). In the present study, induction of ACLY expression by testosterone not only decreased cycling in the TCA cycle, but also promoted glucose assimilation, which played an important role in fatty acid synthesis. Thus, in tissues where fatty acid synthesis is promoted, either β-oxidation-dependent ATP production in mitochondria or the secretion of fatty acids occurs to maintain cellular function (51–53), providing evidence to explain why the presence of fatty acids in the cell has no adverse effect. In accessory reproductive glands, such as the prostate and seminal vesicles, the glucose uptake is activated in a testosterone dependent. However, mitochondrial respiration is not enhanced, cell proliferation is not induced, and the metabolites are used as a substrate for sperm motility.

Revised version: Testosterone has been reported to cause hypertrophy of cardiac myocytes (44, 45). Cell hypertrophy is partly due to the accumulation of fatty acids in the cytoplasm, which causes significant damage to cellular functions (46, 47). On the other hand, normal cell hypertrophy, especially in myofibers, involves the activation of signaling molecules such as Myc, hypoxia-inducible factor, and Pi3k-Akt-mTor, which regulate metabolic reprogramming in cancer by increasing glucose uptake in the presence of oxygen (48–50). In the present study, induction of ACLY expression by testosterone not only decreased cycling in the TCA cycle, but also promoted glucose assimilation, which played an important role in fatty acid synthesis. Thus, in tissues where fatty acid synthesis is promoted, either β-oxidation-dependent ATP production in mitochondria or the secretion of fatty acids occurs to maintain cellular function (51–53) providing evidence to explain why the presence of fatty acids in the cell has no adverse effect. In accessory reproductive glands, such as the prostate and seminal vesicles, the glucose uptake is activated in a testosterone dependent. However, mitochondrial respiration is not enhanced, cell proliferation is not induced, and the metabolites are used as a substrate for sperm motility. This may be because, in the prostate, pyruvate synthesized by the glycolytic system in response to testosterone is converted to acetyl CoA, which is then converted to citrate in the TCA cycle, citrate is released as the main component of seminal plasma (37).

Removed in the revised version: This may be because, in the prostate, pyruvate synthesized by the glycolytic system in response to testosterone is converted to acetyl CoA, which is then converted to citric acid in the TCA cycle, citric acid is released as the main component of seminal plasma (37).

Earlier version: The testosterone-dependent growth of cancerous prostate epithelial cells may be due to genetic mutations, such as elevated expression of the gene encoding aconitase, which converts citric acid to isocitrate, which blocks the TCA cycle and inhibits citric acid secretion (54). On the other hand, in the analysis of seminal epithelial cells, testosterone treatment induced the expression of Acc, which encodes the rate-limiting enzyme for fatty acid synthesis, and Elovl6, which desaturates to oleic acid, significantly increasing total saturated fatty acid and oleic acid content. Therefore, testosterone promotes glucose conversion into fatty acids, especially oleic acid, in seminal epithelial cells. Although there is a commonality in the activation of the glycolytic system testosterone target tissues, the different expression patterns of genes involved in metabolism in mitochondria may exert their specific functions in each tissue. In addition to genes involved in fatty acid synthesis such as Fasn and Elovl6, Acly is also encoded in a region in close proximity to human chromosome 17 and mouse chromosome 11, suggesting that in seminal vesicle epithelial cells, this region is targeted to promote gene expression. The co-regulation of these genes in seminal vesicle epithelial cells leads to synthesizing oleic acid that ensures sperm motility and function.

Revised version: The testosterone-dependent growth of cancerous prostate epithelial cells may be due to genetic mutations, such as elevated expression of the gene encoding aconitase, which converts citrate to isocitrate, which blocks the TCA cycle and inhibits citrate secretion (54). On the other hand, in the analysis of seminal vesicle epithelial cells, testosterone treatment induced the expression of Acc, which encodes the rate-limiting enzyme for fatty acid synthesis, and Elovl6, which desaturates to oleic acid, significantly increasing total saturated fatty acid and oleic acid content. Therefore, testosterone promotes glucose conversion into fatty acids, especially oleic acid, in seminal vesicle epithelial cells. Although there is a commonality in the activation of the glycolytic system testosterone target tissues, the different expression patterns of genes involved in metabolism in mitochondria may exert their specific functions in each tissue. In addition to genes involved in fatty acid synthesis such as Fasn and Elovl6, Acly is also encoded in a region in close proximity to human chromosome 17 and mouse chromosome 11, suggesting that in seminal vesicle epithelial cells, this region is targeted to promote gene expression. The co-regulation of these genes in seminal vesicle epithelial cells leads to synthesizing the fatty acids. Decreased testosterone synthesis in the testes is a phenomenon observed with aging associated with decreased spermatogenesis(55). However, the probability of obtaining sperm without ICSI or IVF is extremely low, despite the presence of sufficient numbers of sperm with normal-morphology (56–58). This indicates that testosterone plays a major role not only in spermatogenesis but also in the functional maturation of sperm. In this study, metabolic abnormalities induced by inhibition of androgen receptor function caused abnormal proliferation of seminal vesicle epithelial cells in aged mice. In other words, the findings of the seminal vesicle epithelial cell culture experiments and flutamide treatment may mimic the changes that occur with aging. Oleic acid synthesis by testosterone-dependent metabolic pathways was regulated in seminal vesicle epithelial cells. It could be assumed that a decrease in testosterone levels causes abnormalities in the components of human semen. In particular, the results open up the possibility of a translational study to detect low testosterone levels in men by testing for oleic acid in semen.

Removed in the revised version: Decreased testosterone synthesis in the testes is a phenomenon observed with aging associated with decreased spermatogenesis(55). However, the probability of obtaining sperm without ICSI or IVF is extremely low, despite the presence of sufficient numbers of sperm with normal-morphology (56–58). This indicates that testosterone plays a major role not only in spermatogenesis but also in the functional maturation of sperm. In this study, metabolic abnormalities induced by inhibition of androgen receptor function caused abnormal proliferation of seminal vesicle epithelial cells in aged mice. In other words, the findings of the seminal vesicle epithelial cell culture experiments and flutamide treatment may mimic the changes that occur with aging. Oleic acid synthesis by testosterone-dependent metabolic pathways was regulated in seminal vesicle epithelial cells and the oleic acid stimulated sperm motility and enhanced in vivo fertilization success. It could be assumed that a decrease in testosterone levels causes abnormalities in the components of human semen. Specifically, the results open the possibility of translational studies to understand human male infertility, especially in cases of failure to conceive by natural mating or artificial insemination.

Earlier version: Because seminal plasma contains components specific to certain male reproductive organs, differences in protein composition may reflect pathological processes in these specific organs (59). However, although proteomic and metabolomic analyses of seminal plasma have been performed (31, 59), the functional changes in seminal plasma, especially the effects on in vivo fertilization potential, are not fully understood. Differences in fatty acid concentrations, including oleic acid, in seminal plasma may provide a predictive marker for the fertilization potential of both human and domestic animal semen. Furthermore, the addition of oleic acid to low-quality semen improved fertilization rates associated with artificial insemination, offering a strategy for improving artificial insemination. In our previous study, the addition of fatty acids to the thawing medium of frozen bull sperm improved linear motility and viability (21). The similarities between mouse and human seminal vesicle epithelial cells demonstrated in this study have strong potential for the clinical application of these observations.

Revised version: Because seminal plasma contains components specific to certain male reproductive organs, differences in protein composition may reflect pathological processes in these specific organs (59). However, although proteomic and metabolomic analyses of seminal plasma have been performed (31, 59), the functional changes in seminal plasma, especially the effects on in vivo fertilization potential, are not fully understood. This is because there are many challenges to measuring the direct relationship between sperm metabolism and movement patterns. It has been suggested that the potential importance of the effect of oleic acid on improving fertility, yet it is too early to conclude this in this study. If a direct causal relationship between fatty acid composition and improved sperm motility is established, differences in the concentration of fatty acids in the seminal plasma will become a predictive marker of the fertilizing ability of semen. The interspecies differences between humans and mice should be interpreted with caution, and the translational aspects to humans are speculative, but this study revealed a testosterone-dependent mechanism of oleic acid production in seminal vesicle epithelium.

Earlier version: Overall, the results of the studies described herein strongly support the hypothesis that testosterone-dependent glucose to oleic acid conversion processes are required for normal seminal plasma functions in the seminal vesicle and normal reproductive performance in males (Fig.10). This study advances our understanding of the role of testosterone in glucose metabolism and fatty acid synthesis and lays the groundwork for future research and potential therapeutic interventions in reproductive biology and fertility.

Revised version: Overall, the results of the studies described herein strongly support the hypothesis that testosterone-dependent glucose to oleic acid conversion processes are required for normal seminal plasma functions in the seminal vesicle (Fig.9). This study advances our understanding of the role of testosterone in glucose metabolism and fatty acid synthesis and lays the groundwork for future research in reproductive biology and fertility.

Earlier version: Metabolic changes in the seminal vesicle epithelial cells by testosterone improve the fertility of sperm.

Revised version: Testosterone-induced metabolic reprogramming in seminal vesicle epithelial cells.

Earlier version: Left part: Testosterone promotes glucose uptake by changes in the localization of GLUT4 in seminal vesicle epithelial cells. ATP-citrate-lyase (ACLY) activated by testosterone promotes fatty acid synthesis by skipping some tricarboxylic acid (TCA) cycle steps. Simultaneously, cell proliferation in the epididymal epithelial cells is suppressed. Testosterone also promotes the synthesis of oleic acid from acetyl-CoA synthesized by ACLY. Right part: Oleic acid, incorporated into the sperm midpiece, increases the oxygen consumption rate (OCR) in sperm mitochondria and increases ATP production. Therefore, oleic acid in seminal plasma is important for male fertility because it improves sperm linear motility and fertilization ability.

Revised version: Testosterone promotes glucose uptake by changes in the localization of GLUT4 in seminal vesicle epithelial cells. ATP-citrate-lyase (ACLY) activated by testosterone promotes fatty acid synthesis by skipping some tricarboxylic acid (TCA) cycle steps. Simultaneously, cell proliferation in the epididymal epithelial cells is suppressed. Testosterone also promotes the synthesis of oleic acid from acetyl-CoA synthesized by ACLY.

Materials and methods

Study design

Earlier version: The objective of this study was to elucidate the mechanisms of sperm plasma synthesis that enhance sperm fertility. Testosterone-dependent changes in the seminal plasma function were examined by comparative analysis of adult mice and flutamide-treated mice. The bioassays of extracts from male accessory glands by epididymal sperm motility analysis were used to examine the seminal plasma components important for sperm fertilization. We also isolated the mouse seminal vesicle epithelial cultured cells, and investigated the synthesis mechanism of the seminal plasma components using RNA-seq, FluxAnalyzer, and shRNA knockdown. Finally, we performed validation using the human epithelial cells.

Revised version: The objective of this study was to elucidate the mechanisms of seminal plasma synthesis. Testosterone-dependent changes in the seminal plasma function were examined by comparative analysis of adult mice and flutamide-treated mice. The bioassays of extracts from male accessory glands by epididymal sperm motility analysis were used to examine the seminal plasma components important for sperm function. We also isolated the mouse seminal vesicle epithelial cultured cells, and investigated the synthesis mechanism of the seminal plasma components using RNA-seq, FluxAnalyzer, and shRNA knockdown. Finally, we performed validation using the human epithelial cells.

Chemicals and Animals

Earlier version: Unless otherwise noted, chemicals used in this study were purchased from Sigma-Aldrich (St. Louis, MO, USA) or Nacalai tesque (Osaka, Japan). Pregnant mare serum gonadotropin (PMSG) and human chorionic gonadotropin (hCG) were purchased from Asuka Seiyaku (Tokyo, Japan). All procedures during animal experiments were reviewed and approved by the Animal Care and Use Committee of Hiroshima University (Hiroshima, Japan; C21-9-3) and conducted according to regulations. Specific pathogen-free C57BL/6NCrl male mice or Crl: CD1 (ICR) female mice were purchased from Jackson Laboratory Japan (Kanagawa, Japan) and housed in an environmentally controlled room with a 12-hour light/dark cycle, a temperature of 23 ± 3°C, and free access to laboratory food (MF; Oriental Yeast Co., Ltd., Tokyo, Japan) and tap water. 8- to 16-week-old male or female mice were used for the experiments. Flutamide-treated mice were injected subcutaneously for 7 days with 50 µg/g (body weight) of flutamide dissolved in corn oil. Vehicle-treated mice received 100 µL of corn oil subcutaneously for 7 days. The aged males were 12 to 16 months old.

Revised version: Unless otherwise noted, chemicals used in this study were purchased from Sigma-Aldrich (St. Louis, MO, USA) or Nacalai tesque (Osaka, Japan). Equine chorionic gonadotropin (eCG) and human chorionic gonadotropin (hCG) were purchased from Asuka Seiyaku (Tokyo, Japan). All procedures during animal experiments were reviewed and approved by the Animal Care and Use Committee of Hiroshima University (Hiroshima, Japan; C21-9-3) and conducted according to regulations. Specific pathogen-free C57BL/6NCrl male mice or Crl: CD1 (ICR) female mice were purchased from Jackson Laboratory Japan (Kanagawa, Japan) and housed in an environmentally controlled room with a 12-hour light/dark cycle, a temperature of 23 ± 3°C, and free access to laboratory food (MF; Oriental Yeast Co., Ltd., Tokyo, Japan) and tap water. 8- to 16-week-old male or female mice were used for the experiments. Flutamide-treated mice were injected subcutaneously for 7 days with 50 µg/g (body weight) of flutamide dissolved in corn oil. Vehicle-treated mice received 100 µL of corn oil subcutaneously for 7 days. The aged males were 12 to 16 months old.

Bioassay using pseudo-seminal plasma and epididymal sperm

Prostate (mixed anterior, dorsal, lateral, and ventral) and seminal vesicles were collected from adult C57BL/6NCrl mice. Seminal vesicle secretions were obtained by pressing the seminal vesicles immediately after euthanasia. Prostatic extracts were obtained by homogenizing the prostate immediately after euthanasia. Seminal vesicle secretions and prostate extracts collected from one male mouse were dissolved in 500 µL and 200 µL of mHTF medium without BSA (60), respectively, then centrifuged at 2000 × g for 5 min at room temperature, and the supernatant was collected, and protein levels were determined. The seminal vesicle secretions from flutamide-treated mice were also collected similarly.

Removed in the revised version: Cauda epididymis sperm recovered from male mouse were dispersed in 500 µL of mHTF without BSA drops and centrifuged at 300 × g, 3 min to wash out epididymis-derived factors. The final protein concentration of prostate and seminal vesicle secretion was adjusted to 10 mg/mL.

Earlier version: An aliquot (2 mL) of the pooled prostate or seminal vesicle extracts was conducted the viscosity measurement data using Tuning Fork Vibro Viscometers (SV-1A, A&D Company, Tokyo, Japan). The epididymis sperm was mixed with seminal vesicle secretion and/or prostate extract as shown in Fig.1 and Supplementary Figure 1. The samples were incubated at 37°C for 60 min in a humidified atmosphere of 5% CO2 in air. Sperm motility patterns were then examined using computer-assisted sperm analysis (CASA; HT CASA-Ceros II, Hamilton Thorne, Beverly, MA, USA), with software version 1.11.9. The following sperm motility parameters were analyzed: percentage of motile sperm (denoted by a path velocity >40 μm/s and a VSL >20%), percentage of motile sperm with progressive motility (denoted by a path velocity >40 μm/s and a straightness ratio >60%). Straight-Line Velocity; VSL: The distance between the first and last tracking points of the sperm trajectory divided by the elapsed time. Curvi-Linear Velocity; VCL: The sum of the distances between the centers of brightness in each frame divided by the elapsed time. Linearity; LIN: Measures the straightness of the path. LIN is derived from the ratio of VSL/VCL multiplied by 100.

Revised version: Cauda epididymis sperm recovered from male mouse were dispersed in 500 µL of mHTF without BSA and centrifuged at 300 × g, 3 min to wash out epididymis-derived factors. The final protein concentration of prostate and seminal vesicle secretion was adjusted to 10 mg/mL. An aliquot (2 mL) of the pooled prostate or seminal vesicle extracts was conducted the viscosity measurement data using Tuning Fork Vibro Viscometers (SV-1A, A&D Company, Tokyo, Japan). The epididymis sperm was mixed with seminal vesicle secretion and/or prostate extract as shown in Fig.1 and Supplementary Figure 1. The samples were incubated at 37°C for 60 min in a humidified atmosphere of 5% CO2 in air. Sperm motility patterns were then examined using computer-assisted sperm analysis (CASA; HT CASA-Ceros II, Hamilton Thorne, Beverly, MA, USA), with software version 1.11.9. The following sperm motility parameters were analyzed: percentage of motile sperm (denoted by a path velocity >40 μm/s and a VSL >20 μm/s), percentage of motile sperm with progressive motility (denoted by a path velocity >40 μm/s and a straightness ratio >60%). Straight-Line Velocity; VSL: The distance between the first and last tracking points of the sperm trajectory divided by the elapsed time. Curvi-Linear Velocity; VCL: The sum of the distances between the centers of brightness in each frame divided by the elapsed time. Linearity; LIN: Measures the straightness of the path. LIN is derived from the ratio of VSL/VCL multiplied by 100.

Serum testosterone levels

The serum was separated from the clot by centrifugation at 1,500 × g for 15 min. Testosterone levels in the serum samples were determined by a rodent testosterone ELISA test kit (ERKR7016, Endocrine Technologies, Inc., Newark, CA, USA) according to the manufacturer’s manual.

Mitochondrial activity

Earlier version: Mitochondrial activity of sperm was measured using a MitoPT® JC-1 Assay Kit (911, Immuno Chemistry Technologies, LLC, Bloomington, MN, USA) according to our previous study (20). Briefly, sperm were incubated with 200 µL of working solution containing 5,5’,6,6’-tetrachloro-1,1’,3,3’-tetraethylbenzimidazolyl carbocyanine iodide (JC-1) dye at 37 °C for 30 min in the dark. Using 10 µM of antimycin control (A8674; Sigma) under the same incubation conditions, we confirmed that the JC-1 reaction is sensitive to changes in membrane potential. The sperm suspension was centrifuged and washed twice with the base medium. After washing, the sperm pellet was resuspended in the base medium and analyzed with an Attune® NxTAcoustic Focusing Cytometer (Thermo Fisher Scientific Inc., Waltham, MA, USA) using a 488 nm laser and a filter with a bandwidth of 574/26 nm. The intensity of the average value was analyzed as the mean fluorescence intensity (MFI) of JC-1 orange aggregates. A total of 100,000 sperm events were analyzed. The sperm population is shown in Supplementary Figure 6.

Revised version: Mitochondrial activity of sperm was measured using a MitoPT® JC-1 Assay Kit (911, Immuno Chemistry Technologies, LLC, Bloomington, MN, USA) according to our previous study (20). Briefly, sperm were incubated with 200 µL of working solution containing 5,5’,6,6’-tetrachloro-1,1’,3,3’-tetraethylbenzimidazolyl carbocyanine iodide (JC-1) dye at 37 °C for 30 min in the dark. Using 10 µM of antimycin control (A8674; Sigma) under the same incubation conditions, we confirmed that the JC-1 reaction is sensitive to changes in membrane potential. The sperm suspension was centrifuged and washed twice with the base medium. After washing, the sperm pellet was resuspended in the base medium and analyzed with an Attune® NxT Acoustic Focusing Cytometer (Thermo Fisher Scientific Inc., Waltham, MA, USA) using a 488 nm laser and a filter with a bandwidth of 574/26 nm. The intensity of the average value was analyzed as the mean fluorescence intensity (MFI) of JC-1 orange aggregates. A total of 100,000 sperm events were analyzed. The sperm population is shown in Supplementary Figure 6.

In vitro fertilization (IVF)

Earlier version: 3 weeks old CD1 female mice were injected intraperitoneally with 4IU eCG to stimulate follicle development and 5 IU hCG at 48 hours later. COCs collected from the oviductal ampulla 16 h after hCG injection were placed into 100 μL of human tubal fluid (HTF) medium. Sperm were collected from the cauda epididymis of 15 weeks or older C57BL/6N male mice in 100 μL of HTF medium. After 60 min of incubation with or without treatments, the sperm were added into HTF medium at a final concentration of 2 × 105 spermatozoa/mL for conventional experiments or 1 × 104 10 nM OA treated sperm/mL, and coincubated with the COCs. After 6h incubation, oocytes were washed three times thoroughly and cultured in 250 μL of KSOM medium (MR 106-D; Sigma). The cleavage rate was determined on Day 1 (Day 0=day of IVF), and the appearance of blastocysts was recorded on Day 5.

Revised version: 3 weeks old CD1 female mice were injected intraperitoneally with 4IU eCG to stimulate follicle development and 5 IU hCG at 48 hours later. COCs collected from the oviductal ampulla 16 h after hCG injection were placed into 100 μL of human tubal fluid (HTF) medium. Sperm were collected from the cauda epididymis of 15 weeks or older C57BL/6N male mice in 100 μL of HTF medium. After 60 min of incubation with or without treatments (Ctrl, Flu or HTF), the sperm were added into HTF medium at a final concentration of 2 × 105 spermatozoa/mL, and coincubated with the COCs. After 6h incubation, oocytes were washed three times thoroughly and cultured in 250 μL of KSOM medium (MR 106-D; Sigma). The cleavage rate was determined on Day 1 (Day 0=day of IVF), and the appearance of blastocysts was recorded on Day 5.

Immunohistochemistry and immunofluorescence

Seminal vesicles were fixed overnight in 4% (w/v) paraformaldehyde/PBS, dehydrated in 70% (v/v) ethanol, and embedded in paraffin. Paraffin-embedded fixed sections (4 μm) were deparaffinized and hydrated with xylene and ethanol, respectively. Some slides were stained with hematoxylin and eosin (HE). For antigen retrieval, sections were placed in citrate buffer (pH 6.0), microwaved until boiling and allowed to stand at the temperature of pre-boil (95°C) for 15 minutes. Sections were incubated with 3% (v/v) hydrogen peroxide in methanol to block endogenous peroxidase and then incubated with Normal Goat Serum Blocking Solution for rabbit antibody (S-1000, Vector Laboratories, Newark, CA, USA) to block nonspecific sites. The sections were then incubated overnight with primary rabbit antibodies: anti-AR antibody (1:100; ab133273; Abcam, Cambridge, UK), anti-Ki67 antibody (1:1000; ab15580; Abcam) or anti-ACLY antibody (1:200; ab40793; Abcam). Positive signals were visualized by DAB (3,3′-Diaminobenzidine, 040-27001, Fujifilm Wako, Osaka, Japan) using the VECTASTAINE LifeABC kit (PK-6101, Vector Laboratories) according to the manufacturer’s protocol. Apoptotic cells (including fragmented DNA) were detected using the In Situ Cell Death Detection Kit POD (11684817910, Roche Ltd., Basel, Switzerland). To avoid false-positive or false-negative results, DNAse-treated sections were used as positive controls, and TdT-untreated sections were used as negative controls for apoptosis.

Nuclei were visualized with hematoxylin (Sakura Finetek, Tokyo, Japan). Ki-67 and apoptotic positive cells were quantified using ImageJ software (National Institutes of Health; Bethesda, MD, USA), with software version 2.1.0/1.53u. Randomly selected sections were stained, and images were obtained from randomly selected fields in each section. The positive percentage is defined as the percentage of positively stained cells in the total number of epithelial cells evaluated. Only positive staining was considered, regardless of staining intensity. For immunofluorescence, after blocking nonspecific sites, the sections were incubated with a primary antibody: anti-GLUT4 antibody (1:200; ab33780; Abcam). After washing with 0.3% (v/v) Triton X-100 in PBS (-), Cy3-labeled goat anti-rabbit IgG (1:200, C-2306, Sigma) and DAPI (VECTESHIELD Mounting Medium with DAPI, H-1200, Vector Laboratories) were used to visualize the antigen and nuclei. Digital images were taken using an APX100 Digital Imaging System and CellSens imaging software (EVIDENT, Tokyo, Japan), with software version 4.1.1.

Isolation and culture of seminal vesicle epithelial cells

Seminal vesicles were removed from euthanized mice, cut open vertically to expose the seminal vesicle cavity, and incubated with Hanks’ balanced salt solution (HBSS; CaCl2・2H2O 0.185 g/L, MgSO4 0.098 g/L, KCl 0.400 g/L, KH2PO4 0.060 g/L, NaHCO3 0.350 g/L, NaCl 8.000 g/L, Na2HPO4 0.048 g/L, D-Glucose 1.0 g/L, dissolved in ultrapure water). Tissues were digested with 10 mL of HBSS containing 25 mg/mL pancreatin and 0.25% (v/v) trypsin EDTA for 1 hour at 4°C, 45 minutes at room temperature, and 15 minutes at 37°C, and epithelial cells were obtained. These were passed through a 70 µm nylon filter (352350, Corning Inc., Corning, NY, USA) and seeded into collagen-coated 12-well plates (4815-010, IWAKI, Shizuoka, Japan). Cells were cultured in DMEM/F-12, 10% fetal bovine serum (FBS; 10438-018, Thermo Fisher Scientific Inc.), 100 U/mL penicillin, 100 µg/mL streptomycin. The culture condition was 37°C, 5% CO2, 95% air, and humidity.

Human seminal vesicle epithelial cells (4460, ScienCell Research Laboratories, Carlsbad, CA, USA) were purchased, seeded in 2 μg/cm2 poly-L-lysine-coated plates, and cultured in epithelial cell medium (4101, ScienCell Research Laboratories). The cells were seeded at 10000 cells/cm2 on collagen-coated plates the day before 100 ng/mL testosterone treatment. Testosterone was dissolved in ethanol, and the final concentration of ethanol was adjusted to 0.1% (v/v) when used for cultured cells. The culture medium was replaced every 48 hours.

Cell proliferation curve

The seminal vesicle epithelial cells were seeded at 10000 cells/cm2 and then were cultured in various concentrations of testosterone. The cells were washed with HBSS (without CaCl2-2H2O/MgSO4) at the time of cell counting and collected using 0.25% (v/v) trypsin. The recovered cells were suspended in DMEM/F-12 medium, and cell numbers were determined using an Automated Cell Counter TC20TM (BIO-RAD Laboratories, Hercules, CA, USA) to calculate the number of cells contained per well (4 cm2).

Cell cycle analysis

The recovered cells as described above were fixed in 2-5 ml of cold (4°C) 70% methanol for at least 30 minutes on ice. After centrifugation at 500g for 5 minutes at 4°C to obtain pellets, the cell pellet was suspended in 50 µM propidium iodide and 20 × diluted RNase (19101, QIAGEN Inc., Venlo, Netherlands) and then were incubated at 37°C for 30 minutes. After twice washing with 1 ml of PBS at 300 g for 3 min at 4°C, the cells were analyzed by flow cytometry. The cell population is shown in Supplementary Figure 8. A total of 20,000 events were analyzed.

RNA-sequencing

Total RNA was extracted from cells using the RNeasy Mini Kit (74106, QIAGEN). Three independent samples were assayed to evaluate the reproducibility of the experimental procedures. The quality of total RNA (1 µg) submitted for sequencing was checked using a Bioanalyzer (Agilent Technologies Inc., Santa Clara, CA, USA), and all RNA integrity numbers were >7.0. The libraries for RNA sequencing were sequenced with 2 × 150-bp paired-end reads on a NovaSeq6000 (Illumina Inc., San Diego, CA, USA). The obtained sequencing data were analyzed using RaNA-seq (61), an open bioinformatics tool for rapid RNA-seq data analysis. FASTQ files were preprocessed in the pipeline using the Fastp tool, and expression was quantified using Salmon. DEseq2 was used to identify differentially expressed genes (DEGs) with adjusted p-values less than 0.05 between the testosterone-treated groups in transcripts per million (TPM) between control and testosterone-treated cells. The pathways were identified from the KEGG database on RaNA-seq.

Extracellular flux analysis

Oxygen consumption rate (OCR) and ECAR were measured using an extracellular flux analyzer (XF HS Mini; Agilent Technologies). Eight-well plates for the XF HS Mini were coated with fibronectin (5 µg/mL, F1141, Sigma-Aldrich) dissolved in phosphate-buffered saline (PBS), incubated overnight, and then washed with Roswell Park Memorial Laboratory medium for XF (Roswell Park Memorial Institute medium; RPMI, 103576-100, Agilent Technologies; containing 1% FBS) for XF before analysis. Mouse seminal vesicle epithelial cells or human seminal vesicle epithelial cells treated with 100 ng/mL testosterone for 7 days were suspended in RPMI medium at a concentration of 10,000 cells/180 µL, or mouse seminal vesicle epithelial cells infected with shRNA were suspended at a concentration of 2000 cells/180 µL in each well. After centrifugation at 300 × g for 3 minutes, the plates were placed in an incubator at 37°C and 100% air and used for analysis within 1 hour.

For glucose metabolism assessment, cells were sequentially treated as indicated with D-Glucose (10 mM; 103577-100, Agilent Technologies), Oligomycin (1 µM; O4876, Sigma-Aldrich), and 2-Deoxy-Glucose (50 mM; D8375, Sigma-Aldrich).

Earlier version: For mitochondrial respiration measurement, cells were sequentially treated as indicated with Oligomycin (1 µM), carbonyl cyanide 4-(trifluoromethoxy) phenylhydrazone (FCCP; 5 µM, SML2959, Sigma-Aldrich), and Rotenone (1 µM; R8875, Sigma-Aldrich)/Antimycin A (1 µM; A8674, Sigma-Aldrich). Cellular metabolism readouts, such as glycolysis, glycolytic reserve, basal OCR, Maximum respiration, ATP Production, Spare respiratory capacity, were determined using the Agilent Seahorse Wave software with software version 3.0.0.41. However, to avoid many ambiguities and problems with interpretation, “glycolysis” is written as “Glucose-induced ECAR”, “glycolytic reserve” as “Oligomycin-sensitive ECAR”, “maximum respiration” as “FCCP uncoupled respiration”, and “ATP production” as “Oligomycin-sensitive respiration”.

Revised version: For mitochondrial respiration measurement, cells were sequentially treated as indicated with Oligomycin (1 µM), carbonyl cyanide 4-(trifluoromethoxy) phenylhydrazone (FCCP; 5 µM, SML2959, Sigma-Aldrich), and Rotenone (1 µM; R8875, Sigma-Aldrich)/Antimycin A (1 µM; A8674, Sigma-Aldrich). Cellular metabolism readouts, such as glycolysis, glycolytic reserve, basal OCR, Maximum respiration, ATP Production, Spare respiratory capacity, were determined using the Agilent Seahorse Wave software with software version 3.0.0.41. However, to avoid many ambiguities and problems with interpretation, “glycolysis” is written as “Glucose-induced

Added in the revised version: ECAR”, “glycolytic reserve” as “Oligomycin-sensitive ECAR”, “maximum respiration” as “FCCP uncoupled respiration”, and “ATP production” as “Oligomycin-sensitive respiration”.

Removed in the revised version: For mouse sperm analysis, NaHCO3-free HTF medium was prepared according to Balbach et al. (62). Plates were coated with concanavalin A (0.5 mg/mL, Fujifilm Wako Chemicals) overnight the day before assay. Each well was seeded with 180 μL of NaHCO3- Free HTF medium containing 3,000,000 sperm cells. Sperm were injected with oleic acid (final concentration: 10 nM) with or without etomoxir (final concentration: 40 nM, E1905, Sigma-Aldrich).

Measurement of pyruvate and lactic acid

For measurements, mouse or human seminal vesicle epithelial cells were treated with 100 ng/ml testosterone for 6 days. The culture supernatant was collected after 24 hr of additional culture. Pyruvate consumption was measured using the Pyruvate Assay kit (MET-5029, Cell Biolabs, Inc., San Diego, CA, USA). The fluorescence intensity was measured using a microplate reader (ex. 570 nm/em. 615 nm). Lactate released into the medium was measured using the Lactate Assay kit (L256, DOJINDO LABOLATORIES Co., Ltd., Kumamoto, Japan). The absorbance was measured using a microplate reader (wave length=450nm).

qPCR

The cDNA was synthesized from 40 ng of total RNA using oligo(dT)15 primers (3805, Takara Bio Inc., Shiga, Japan) and an Avian myeloblastosis virus (AMV) reverse transcriptase from Promega (M5101, Madison, WI, USA). The cDNA and primers were added to Power SYBR Green PCR Master Mix (4367659, Applied Biosystems, Foster City, CA, USA) in a total reaction volume of 15 μl. Conditions were set to the following parameters: 10 min at 95°C, 15 sec at 95°C, and 1 min at 60°C. Cycles were repeated 45 times. Specific primer pairs were selected and analyzed, as shown in Table S1. Expression was first normalized to housekeeping gene Rpl19, and fold change was calculated relative to the mean of the control samples.

Glucose uptake test of seminal vesicle epithelial cells and lipid determination of culture supernatants

Mouse or human seminal vesicle epithelial cells treated with or without 100 ng/mL testosterone for 6 days were used for this experiment. Cells were treated with or without the GLUT4 inhibitor indinavir (100 nM) and 100 ng/mL testosterone for an additional 24 hours. Collected 15,000 cells were subjected to Glucose Uptake Assay Kit-Green (UP02; DOJINDO), and fluorescence intensity was measured using a microplate reader (ex. 485 nm/em. 535 nm).

Total lipid content in cultured medium was quantified using a lipid quantification kit (STA-617 Fluorometric, Cell Biolabs, Inc.) according to the manufacturer’s protocol. The fluorescence intensity was measured using a microplate reader (ex. 485 nm/em. 535 nm).

Gas chromatography-mass spectrometry analysis

Fatty acid measurements in culture supernatant and seminal vesicle secretions were done according to our previous work (21). Fatty acids in the samples were extracted and methylated using a fatty acid methylation kit (06482-04, Nacalai Tesque) and then purified using a fatty acid methyl ester purification kit (06483-94, Nacalai Tesque). The methylated sample (3.0 mL) was dried and dissolved in 50 μL of elution solution. Samples were subjected to GC-MS.

Fatty acids were determined on an Agilent 7890A GC system coupled with a JMS-T 100 GCv mass detector at the Natural Science Research Center, Hiroshima University (N-BARD). Sample (1.0 µL aliquot) was injected to SP®-2560 capillary GC column (100 m × 0.25 mm i.d. × 0.2 μm film thickness; 24056, Sigma-Aldrich). For measurements, helium was used as the carrier gas at a constant flow rate of 2.0 mL/min. The GC oven temperature was gradually increased from 100°C to 300°C at a rate of 10°C/min and held at 300°C for 20 minutes. Ionization was performed using electron ionization with an electron energy of 70 eV. Mass spectra were obtained in scan mode (mass scan range 29-800 m/z). NIST MS Search v.2.0 was used to detect and identify of fatty acids.

Added in the revised version: Sperm motility parameters with the lipid Mixture

Added in the revised version: Sperm recovered from the cauda epididymis of male mice were dispersed in 500 µL of mHTF without BSA and centrifuged at 300 × g for 3 minutes to wash out epididymis-derived factors. The sperm pellets were resuspended in HTF medium with or without a lipid mixture (LM) that included 2 μg/mL arachidonic acid, 10 μg/mL linoleic acid, 10 μg/mL linolenic acid, 10 μg/mL myristic acid, 10 μg/mL oleic acid, 10 μg/mL palmitic acid, and 10 μg/mL stearic acid (L0288, Sigma-Aldrich). The sperm were then incubated for 60 minutes at 37°C under a humidified atmosphere of 5% CO2. Sperm motility patterns were subsequently examined using CASA.

Removed in the revised version: Quantification of sperm oleic acid uptake by flow cytometry

Removed in the revised version: To detect the fluorescence intensity of each sperm sample using flow cytometry, sperm collected from the epididymis were incubated with 1, 10 or 100 μM of fluorescently labeled oleic acid (F-OA-01, NANOCS, New York, NY, USA). After washing, fluorescence were analyzed by flow cytometry using a 530/30 nm bandwidth filter and measured as mean intensity. A total of 100,000 sperm events were analyzed. The localization of fluorescence in sperm was determined using an APX100 Digital Imaging System.

Removed in the revised version: Evaluation of acrosome damage

Removed in the revised version: The acrosome status of sperm was assessed using 25 µg/mL fluorescein isothiocyanate- conjugated peanut agglutinin (FITC-PNA; L7381, Sigma Aldrich, St. Louis, MO) and 2.4 nM propidium iodide (PI; L 7011, Invitrogen) were used for evaluation. After incubation in the dark at 37°C for 10 min. The fluorescence of FITC and PI were analyzed by flow cytometry using a 530/30 and 695/40 nm bandwidth filter, respectively. Damaged sperm were analyzed—FITC fluorescence positive, but PI negative gate (Q4). A total of 100,000 sperm events were analyzed. The sperm population is shown in Supplementary Figures 5 and 6.

ACLY knock down experiment using shRNA

The U6 shRNA lentiviral vector of the Acly gene (mAcly[shRNA#1]; a virus packaging service, VectorBuilder, Chicago, IL, USA) was used for knockdown experiments. The ACLY sequence targeted by the shRNA vector used in this study was 5′-GCAGCAAAGATGTTCAGTAAA-3′. Lentivirally transduced cells were prepared according to manufacture protocol. Immediately after isolation, the mouse seminal vesicle epithelial cells were cultured overnight in a medium containing 5 µg/mL polybrene (VectorBuilder) and lentiviral particles expressing an shRNA targeting ACLY or a scramble (Multiplicity of Infection of 10). After an additional overnight incubation in a fresh culture medium, the cells were incubated for 5 days in a medium containing 1 µg/ml puromycin (19752-64, Nacalai tesque). 100 ng/mL of testosterone was added to all procedures for the testosterone-treated group. Infection efficiency, calculated based on the number of GFP-positive cells using an APX100 Digital Imaging System, ranged between 70% and 90%.

Added in the revised version: For sperm motility analysis, sperm were collected and washed with mHTF from the culture supernatant and cauda epididymis. The samples were incubated at 37°C for 60 minutes in a humidified atmosphere of 5% CO2 in air. The sperm motility pattern was then examined using CASA.

Western blot analysis

Earlier version: Protein samples of mouse seminal vesicle epithelial cells, seminal vesicle tissue, and sperm were prepared by homogenization in cell lysis buffer (04719964001, cOmplete™ Lysis-M EDTA-free, Roche) and diluted with an equal volume of sample buffer solution containing 2-ME (2 ×) for SDS-PAGE (30566-22, Nacalai tesque). The extracts (10 μg protein) were separated by SDS-polyacrylamide gel (10%) electrophoresis and transferred to polyvinylidene fluoride membranes (10600069, Cytiva, Tokyo, Japan). The membranes were then treated with Tris- buffered saline and Tween 20 (TBST, 20 mM Tris, pH 7.5, 150 mM NaCl, 0.1% Tween 20) containing 5% (w/v) Non-Fat Dry milk (MORINAGA MILK INDUSTRY Co., Ltd., Tokyo, Japan) for the nonspecific reaction was blocked. When detecting phosphotyrosine, containing 5% (w/v) bovine serum albumin was used as a blocking agent. The membranes were incubated overnight at 4°C with primary antibodies: anti-phosphotyrosine antibody (1:10,000; 8959; Cell Signaling), anti-GLUT4 antibody (1:100; ab33780; Abcam), anti-ACLY antibody (1:10,000; ab40793; Abcam) or anti-α/β-Tubulin antibody (1:1,000; 2148S; Cell Signaling). The next day, membranes were washed with TBST and incubated with HRP-labeled antibody specific for rabbit IgG (1:4000; 7074S; Cell Signaling). The bands were visualized by using Enhanched Chemiluminescence detection systems (RPN2232, Cytiva) and ChemiDocTM MP Imaging System (BIO-RAD Laboratories). The bands were analyzed by using the software ImageJ. Densitometry analysis was performed using the ImageJ Gel Analysis tool, and the background of the gel was also removed individually for each band. The full-length blotting images are shown in Supplementary Figure 7.

Revised version: Protein samples of mouse seminal vesicle epithelial cells, seminal vesicle tissue, and sperm were prepared by homogenization in cell lysis buffer (04719964001, cOmplete™ Lysis-M EDTA-free, Roche) and diluted with an equal volume of sample buffer solution containing 2-ME (2 ×) for SDS-PAGE (30566-22, Nacalai tesque). The extracts (10 μg protein) were separated by SDS-polyacrylamide gel (10%) electrophoresis and transferred to polyvinylidene fluoride membranes (10600069, Cytiva, Tokyo, Japan). The membranes were then treated with Tris-buffered saline and Tween 20 (TBST, 20 mM Tris, pH 7.5, 150 mM NaCl, 0.1% Tween 20) containing 5% (w/v) Non-Fat Dry milk (MORINAGA MILK INDUSTRY Co., Ltd., Tokyo, Japan) for the nonspecific reaction was blocked. The membranes were incubated overnight at 4°C with primary antibodies: anti-GLUT4 antibody (1:100; ab33780; Abcam), anti-ACLY antibody (1:10,000; ab40793; Abcam) or anti-α/β-Tubulin antibody (1:1,000; 2148S; Cell Signaling). The next day, membranes were washed with TBST and incubated with HRP-labeled antibody specific for rabbit IgG (1:4000; 7074S; Cell Signaling). The bands were visualized by using Enhanced Chemiluminescence detection systems (RPN2232, Cytiva) and ChemiDocTM MP Imaging System (BIO-RAD Laboratories). The bands were analyzed by using the software ImageJ. Densitometry analysis was performed using the ImageJ Gel Analysis tool, and the background of the gel was also removed individually for each band. The full-length blotting images are shown in Supplementary Figure 7.

Removed in the revised version: ATP production in sperm

Removed in the revised version: After 1 hour of incubation in the presence or absence of oleic acid (10 nM), ATP in sperm was visualized by BioTracker ATP-Red (SCT045, Sigma-Aldrich). The sperm were incubated with 10 nM of working solution containing BioTracker ATP-Red dye at 37 °C for 60 min in the dark. After washing, the sperm pellet was resuspended in the base medium and analyzed with the flow cytometry using a 488 nm laser and a filter with a bandwidth of 574/26 nm. A total of 100,000 sperm events were analyzed. The sperm population is shown in Supplementary Figure 6.

Removed in the revised version: Artificial insemination

Removed in the revised version: To examine the involvement of seminal vesicle secretions in fertilization, sperm from C57BL/6N mice were injected directly into the uterine bodies of superovulated CD1 mice by syringe, as previously described with minor modifications (17, 22). CD1 female mice were injected intraperitoneally with 8 IU eCG to stimulate follicle development and used as recipients for artificial insemination 48 hours later. Artificial seminal fluid was prepared in HTF medium to a final concentration of 10 mg/mL of collected control or frutamide-derived seminal vesicle secretions. As rescue experiment, 10 nM of oleic acid was added to the frutamide-derived seminal secretions. Cauda epididymis sperm was suspended with above artificial seminal fluid (1 × 106 sperm/mL), and a total of 40 µL (20 µL each per uterus) was injected through the cervix of the female mice. After artificial insemination, 7.5 IU of hCG was injected to stimulate ovulation. Subsequently, a mixture of vaseline and paraffin with a ratio of 10:1 was injected to form an artificial plug around the cervix. After 36∼48 hours of sperm injection, the putative eggs were retrieved from the oviducts. The number of two-cell embryos and mature unfertilized eggs were counted and used to estimate the fertilization success rate.

Statistical analysis

Quantitative data were presented as means ± SEM. The Bartlett’s test was used to confirm whether the variances were equal for each group in experiments consisting of three or more groups. All statistical details of the experiments are given in the figure legends. R [(version 4.3.1 (2023-06-16)] was used for statistical analysis. Data were considered statistically significant at p < 0.05.